A Myelin Map of Trunk Folds in the Elephant Trigeminal Nucleus

Author Response:

The following is the authors’ response to the previous reviews.

We carefully read through the second-round reviews and the additional reviews. To us, the review process is somewhat unusual and very much dominated by referee 2, who aggressively insists that we mixed up the trigeminal nucleus and inferior olive and that as a consequence our results are meaningless. We think the stance of referee 2 and the focus on one single issue (the alleged mix-up of trigeminal nucleus and inferior olive) is somewhat unfortunate, leaves out much of our findings and we debated at length on how to deal with further revisions. In the end, we decided to again give priority to addressing the criticism of referees 2, because it is hard to go on with a heavily attacked paper without resolving the matter at stake. The following is a summary of, what we did:

Additional experimental work:

(1) We checked if the peripherin-antibody indeed reliably identifies climbing fibers.

To this end, we sectioned the elephant cerebellum and stained sections with the peripherin-antibody. We find: (i) the cerebellar white matter is strongly reactive for peripherin-antibodies, (ii) cerebellar peripherin-antibody staining of has an axonal appearance. (iii) Cerebellar Purkinje cell somata appear to be ensheated by peripherin-antibody staining. (iv) We observed that the peripherin-antibody reactivity gradually decreases from Purkinje cell somata to the pia in the cerebellar molecular layer. This work is shown in our revised Figure 2. All these four features align with the distribution of climbing fibers (which arrive through the white matter, are axons, ensheat Purkinje cell somata, and innervate Purkinje cell proximally not reaching the pia). In line with previous work, which showed similar cerebellar staining patterns in several species (Errante et al. 1998), we conclude that elephant climbing fibers are strongly reactive for peripherin-antibodies.

(2) We delineated the elephant olivo-cerebellar tract.

The strong peripherin-antibody reactivity of elephant climbing fibers enabled us to delineate the elephant olivo-cerebellar tract. We find the elephant olivo-cerebellar tract is a strongly peripherin-antibody reactive, well-delineated fiber tract several millimeters wide and about a centimeter in height. The unstained olivo-cerebellar tract has a greyish appearance. In the anterior regions of the olivo-cerebellar tract, we find that peripherin-antibody reactive fibers run in the dorsolateral brainstem and approach the cerebellar peduncle, where the tract gradually diminishes in size, presumably because climbing fibers discharge into the peduncle. Indeed, peripherin-antibody reactive fibers can be seen entering the cerebellar peduncle. Towards the posterior end of the peduncle, the olivo-cerebellar disappears (in the dorsal brainstem directly below the peduncle. We note that the olivo-cerebellar tract was referred to as the spinal trigeminal tract by Maseko et al. 2013. We think the tract in question cannot be the spinal trigeminal tract for two reasons: (i) This tract is the sole brainstem source of peripherin-positive climbing fibers entering the peduncle/ the cerebellum; this is the defining characteristic of the olivo-cerebellar tract. (ii) The tract in question is much smaller than the trigeminal nerve, disappears posterior to where the trigeminal nerve enters the brainstem (see below), and has no continuity with the trigeminal nerve; the continuity with the trigeminal nerve is the defining characteristic of the spinal trigeminal tract, however.

The anterior regions of the elephant olivo-cerebellar tract are similar to the anterior regions of olivo-cerebellar tract of other mammals in its dorsolateral position and the relation to the cerebellar peduncle. In its more posterior parts, the elephant olivo-cerebellar tract continues for a long distance (~1.5 cm) in roughly the same dorsolateral position and enters the serrated nucleus that we previously identified as the elephant inferior olive. The more posterior parts of the elephant olivo-cerebellar tract therefore differ from the more posterior parts of the olivo-cerebellar tract of other mammals, which follows a ventromedial trajectory towards a ventromedially situated inferior olive. The implication of our delineation of the elephant olivo-cerebellar tract is that we correctly identified the elephant inferior olive.

(3) An in-depth analysis of peripherin-antibody reactivity also indicates that the trigeminal nucleus receives no climbing fiber input.

We also studied the peripherin-antibody reactivity in and around the trigeminal nucleus. We had also noted in the previous submission that the trigeminal nucleus is weakly positive for peripherin, but that the staining pattern is uniform and not the type of axon bundle pattern that is seen in the inferior olive of other mammals. To us, this observation already argued against the presence of climbing fibers in the trigeminal nucleus. We also noted that the myelin stripes of the trigeminal nucleus were peripherin-antibody-negative. In the context of our olivo-cerebellar tract tracing we now also scrutinized the surroundings of the trigeminal nucleus for peripherin-antibody reactivity. We find that the ventral brainstem surrounding the trigeminal nucleus is devoid of peripherin-antibody reactivity. Accordingly, no climbing fibers, (which we have shown to be strongly peripherin-antibody-positive, see our point 1) arrive at the trigeminal nucleus. The absence of climbing fiber input indicates that previous work that identified the (trigeminal) nucleus as the inferior olive (Maseko et al 2013) is unlikely to be correct.

(4) We characterized the entry of the trigeminal nerve into the elephant brain.

To better understand how trigeminal information enters the elephant’s brain, we characterized the entry of the trigeminal nerve. This analysis indicated to us that the trigeminal nerve is not continuous with the olivo-cerebellar tract (the spinal trigeminal tract of Maseko et al. 2013) as previously claimed by Maseko et al. 2013. We show some of this evidence in Referee-Figure 1 below. The reason we think the trigeminal nerve is discontinuous with the olivo-cerebellar tract is the size discrepancy between the two structures. We first show this for the tracing data of Maseko et al. 2013. In the Maseko et al. 2013 data the trigeminal nerve (Referee-Figure 1A, their plate Y) has 3-4 times the diameter of the olivocerebellar tract (the alleged spinal trigeminal tract, Referee-Figure 1B, their plate Z). Note that most if not all trigeminal fibers are thought to continue from the nerve into the trigeminal tract (see our rat data below). We plotted the diameter of the trigeminal nerve and diameter of the olivo-cerebellar (the spinal trigeminal tract according to Maseko et al. 2013) from the Maseko et al. 2013 data (Referee-Figure 1C) and we found that the olivocerebellar tract has a fairly consistent diameter (46 ± 9 mm2, mean ± SD). Statistical considerations and anatomical evidence suggest that the tracing of the trigeminal nerve into the olivo-cerebellar (the spinal trigeminal tract according to Maseko et al. 2013) is almost certainly wrong. The most anterior point of the alleged spinal trigeminal tract has a diameter of 51 mm2 which is more than 15 standard deviations different from the most posterior diameter (194 mm2) of the trigeminal tract. For this assignment to be correct three-quarters of trigeminal nerve fibers would have to spontaneously disappear, something that does not happen in the brain. We also made similar observations in the African elephant Bibi, where the trigeminal nerve (Referee-Figure 1D) is much larger in diameter than the olivocerebellar tract (Referee-Figure 1E). We could also show that the olivocerebellar tract disappears into the peduncle posterior to where the trigeminal nerve enters (Referee-Figure 1F). Our data are very similar to Maseko et al. indicating that their outlining of structures was done correctly. What appears to have been oversimplified, is the assignment of structures as continuous. We also quantified the diameter of the trigeminal nerve and the spinal trigeminal tract in rats (from the Paxinos & Watson atlas; Referee-Figure 1D); as expected we found the trigeminal nerve and spinal trigeminal tract diameters are essentially continuous.

In our hands, the trigeminal nerve does not continue into a well-defined tract that could be traced after its entry. In this regard, it differs both from the olivo-cerebellar tract of the elephant or the spinal trigeminal tract of the rodent, both of which are well delineated. We think the absence of a well-delineated spinal trigeminal tract in elephants might have contributed to the putative tracing error highlighted in our Referee-Figure 1A-C.

We conclude that a size mismatch indicates trigeminal fibers do not run in the olivo-cerebellar tract (the spinal trigeminal tract according to Maseko et al. 2013).

The trigeminal nerve is discontinuous with the olivo-cerebellar tract (the spinal trigeminal tract according to Maseko et al.

2013). A, Trigeminal nerve (orange) in the brain of African elephant LAX as delineated by Maseko et al. 2013 (coronal section; their plate Y). B, Most anterior appearance of the spinal trigeminal tract of Maseko et al. 2013 (blue; coronal section; their plate Z). Note the much smaller diameter of the spinal trigeminal tract compared to the trigeminal nerve shown in C, which argues against the continuity of the two structures. Indeed, our peripherin-antibody staining showed that the spinal trigeminal tract of Maseko corresponds to the olivo-cerebellar tract and is discontinuous with the trigeminal nerve. C, Plot of the trigeminal nerve and olivo-cerebellar tracts (the spinal trigeminal tract according to Maseko et al. 2013) diameter along the anterior-posterior axis. The trigeminal nerve is much larger in diameter than the olivocerebellar tract (the spinal trigeminal tract according to Maseko et al. 2013). C, D measurements, for which sections are shown in panels C and D respectively. The olivocerebellar tract (the spinal trigeminal tract according to Maseko et al. 2013) has a consistent diameter; data replotted from Maseko et al. 2013. At mm 25 the inferior olive appears. D, Trigeminal nerve entry in the brain of African elephant Bibi; our data, coronal section, the trigeminal nerve is outlined in orange, note the large diameter. E, Most anterior appearance of the olivo-cerebellar tract in the brain of African elephant Bibi; our data, coronal section, approximately 3 mm posterior to the section shown in A, the olivocerebellar tract is outlined in blue. Note the smaller diameter of the olivo-cerebellar tract compared to the trigeminal nerve, which argues against the continuity of the two structures. F, Plot of the trigeminal nerve and olivo-cerebellar tract diameter along the anterior-posterior axis. The nerve and olivo-cerebellar tract are discontinuous and the trigeminal nerve is much larger in diameter than the olivocerebellar tract (the spinal trigeminal tract according to Maseko et al. 2013); our data. D, E measurements, for which sections are shown in panels D and E respectively. At mm 27 the inferior olive appears. G, In the rat the trigeminal nerve is continuous in size with the spinal trigeminal tract. Data replotted from Paxinos and Watson.

Reviewer 2 (Public Review):

As indicated in my previous review of this manuscript (see above), it is my opinion that the authors have misidentified, and indeed switched, the inferior olivary nuclear complex (IO) and the trigeminal nuclear complex (Vsens). It is this specific point only that I will address in this second review, as this is the crucial aspect of this paper - if the identification of these nuclear complexes in the elephant brainstem by the authors is incorrect, the remainder of the paper does not have any scientific validity.

Comment: We agree with the referee that it is most important to sort out, the inferior olivary nuclear complex (IO) and the trigeminal nuclear complex, respectively.Change: We did additional experimental work to resolve this matter as detailed at the beginning of our response. Specifically, we ascertained that elephant climbing fibers are strongly peripherin-positive. Based on elephant climbing fiber peripherin-reactivity we delineated the elephant olivo-cerebellar tract. We find that the olivo-cerebellar connects to the structure we refer to as inferior olive to the cerebellum (the referee refers to this structure as the trigeminal nuclear complex). We also found that the trigeminal nucleus (the structure the referee refers to as inferior olive) appears to receive no climbing fibers. We provide indications that the tracing of the trigeminal nerve into the olivo-cerebellar tract by Maseko et al. 2023 was erroneous (Author response image 1). These novel findings support our ideas but are very difficult to reconcile with the referee’s partitioning scheme.

The authors, in their response to my initial review, claim that I "bend" the comparative evidence against them. They further claim that as all other mammalian species exhibit a "serrated" appearance of the inferior olive, and as the elephant does not exhibit this appearance, that what was previously identified as the inferior olive is actually the trigeminal nucleus and vice versa.

For convenience, I will refer to IOM and VsensM as the identification of these structures according to Maseko et al (2013) and other authors and will use IOR and VsensR to refer to the identification forwarded in the study under review.

The IOM/VsensR certainly does not have a serrated appearance in elephants. Indeed, from the plates supplied by the authors in response (Referee Fig. 2), the cytochrome oxidase image supplied and the image from Maseko et al (2013) shows a very similar appearance. There is no doubt that the authors are identifying structures that closely correspond to those provided by Maseko et al (2013). It is solely a contrast in what these nuclear complexes are called and the functional sequelae of the identification of these complexes (are they related to the trunk sensation or movement controlled by the cerebellum?) that is under debate.

Elephants are part of the Afrotheria, thus the most relevant comparative data to resolve this issue will be the identification of these nuclei in other Afrotherian species. Below I provide images of these nuclear complexes, labelled in the standard nomenclature, across several Afrotherian species.

(A) Lesser hedgehog tenrec (Echinops telfairi)

Tenrecs brains are the most intensively studied of the Afrotherian brains, these extensive neuroanatomical studies undertaken primarily by Heinz Künzle. Below I append images (coronal sections stained with cresol violet) of the IO and Vsens (labelled in the standard mammalian manner) in the lesser hedgehog tenrec. It should be clear that the inferior olive is located in the ventral midline of the rostral medulla oblongata (just like the rat) and that this nucleus is not distinctly serrated. The Vsens is located in the lateral aspect of the medulla skirted laterally by the spinal trigeminal tract (Sp5). These images and the labels indicating structures correlate precisely with that provide by Künzle [(1997, 10.1016](callto:(1997,%2010.1016)/S0168- 0102(97)00034-5), see his Figure 1K,L. Thus, in the first case of a related species, there is no serrated appearance of the inferior olive, the location of the inferior olive is confirmed through connectivity with the superior colliculus (a standard connection in mammals) by Künzle (1997), and the location of Vsens is what is considered to be typical for mammals. This is in agreement with the authors, as they propose that ONLY the elephants show the variations they report.

(B) Giant otter shrew (Potomogale velox)

The otter shrews are close relatives of the Tenrecs. Below I append images of cresyl violet (left column) and myelin (right column) stained coronal sections through the brainstem with the IO, Vsens and Sp5 labelled as per standard mammalian anatomy. Here we see hints of the serration of the IO as defined by the authors, but we also see many myelin stripes across the IO. Vsens is located laterally and skirted by the Sp5. This is in agreement with the authors, as they propose that ONLY the elephants show the variations they report.

(C) Four-toed sengi (Petrodromus tetradactylus)

The sengis are close relatives of the Tenrecs and otter shrews, these three groups being part of the Afroinsectiphilia, a distinct branch of the Afrotheria. Below I append images of cresyl violet (left column) and myelin (right column) stained coronal sections through the brainstem with the IO, Vsens and Sp5 labelled as per standard mammalian anatomy. Here we see vague hints of the serration of the IO (as defined by the authors), and we also see many myelin stripes across the IO. Vsens is located laterally and skirted by the Sp5. This is in agreement with the authors, as they propose that ONLY the elephants show the variations they report.

(D) Rock hyrax (Procavia capensis)

The hyraxes, along with the sirens and elephants form the Paenungulata branch of the Afrotheria. Below I append images of cresyl violet (left column) and myelin (right column) stained coronal sections through the brainstem with the IO, Vsens and Sp5 labelled as per the standard mammalian anatomy. Here we see hints of the serration of the IO (as defined by the authors), but we also see evidence of a more "bulbous" appearance of subnuclei of the IO (particularly the principal nucleus), and we also see many myelin stripes across the IO. Vsens is located laterally and skirted by the Sp5. This is in agreement with the authors, as they propose that ONLY the elephants show the variations they report.

(E) West Indian manatee (Trichechus manatus)

The sirens are the closest extant relatives of the elephants in the Afrotheria. Below I append images of cresyl violet (top) and myelin (bottom) stained coronal sections (taken from the University of Wisconsin-Madison Brain Collection, https://brainmuseum.org, and while quite low in magnification they do reveal the structures under debate) through the brainstem with the IO, Vsens and Sp5 labelled as per standard mammalian anatomy. Here we see the serration of the IO (as defined by the authors). Vsens is located laterally and skirted by the Sp5. This is in agreement with the authors, as they propose that ONLY the elephants show the variations they report.

These comparisons and the structural identification, with which the authors agree as they only distinguish the elephants from the other Afrotheria, demonstrate that the appearance of the IO can be quite variable across mammalian species, including those with a close phylogenetic affinity to the elephants. Not all mammal species possess a "serrated" appearance of the IO. Thus, it is more than just theoretically possible that the IO of the elephant appears as described prior to this study.

So what about elephants? Below I append a series of images from coronal sections through the African elephant brainstem stained for Nissl, myelin, and immunostained for calretinin. These sections are labelled according to standard mammalian nomenclature. In these complete sections of the elephant brainstem, we do not see a serrated appearance of the IOM (as described previously and in the current study by the authors). Rather the principal nucleus of the IOM appears to be bulbous in nature. In the current study, no image of myelin staining in the IOM/VsensR is provided by the authors. However, in the images I provide, we do see the reported myelin stripes in all stains - agreement between the authors and reviewer on this point. The higher magnification image to the bottom left of the plate shows one of the IOM/VsensR myelin stripes immunostained for calretinin, and within the myelin stripes axons immunopositive for calretinin are seen (labelled with an arrow). The climbing fibres of the elephant cerebellar cortex are similarly calretinin immunopositive (10.1159/000345565). In contrast, although not shown at high magnification, the fibres forming the Sp5 in the elephant (in the Maseko description, unnamed in the description of the authors) show no immunoreactivity to calretinin.

Comment: We appreciate the referee’s additional comments. We concede the possibility that some relatives of elephants have a less serrated inferior olive than most other mammals. We maintain, however, that the elephant inferior olive (our Figure 1J) has the serrated appearance seen in the vast majority of mammals.

Change: None.

Peripherin Immunostaining

In their revised manuscript the authors present immunostaining of peripherin in the elephant brainstem. This is an important addition (although it does replace the only staining of myelin provided by the authors which is unusual as the word myelin is in the title of the paper) as peripherin is known to specifically label peripheral nerves. In addition, as pointed out by the authors, peripherin also immunostains climbing fibres (Errante et al., 1998). The understanding of this staining is important in determining the identification of the IO and Vsens in the elephant, although it is not ideal for this task as there is some ambiguity. Errante and colleagues (1998; Fig. 1) show that climbing fibres are peripherin-immunopositive in the rat. But what the authors do not evaluate is the extensive peripherin staining in the rat Sp5 in the same paper (Errante et al, 1998, Fig. 2). The image provided by the authors of their peripherin immunostaining (their new Figure 2) shows what I would call the Sp5 of the elephant to be strongly peripherin immunoreactive, just like the rat shown in Errant et al (1998), and more over in the precise position of the rat Sp5! This makes sense as this is where the axons subserving the "extraordinary" tactile sensitivity of the elephant trunk would be found (in the standard model of mammalian brainstem anatomy). Interestingly, the peripherin immunostaining in the elephant is clearly lamellated...this coincides precisely with the description of the trigeminal sensory nuclei in the elephant by Maskeo et al (2013) as pointed out by the authors in their rebuttal. Errante et al (1998) also point out peripherin immunostaining in the inferior olive, but according to the authors this is only "weakly present" in the elephant IOM/VsensR. This latter point is crucial. Surely if the elephant has an extraordinary sensory innervation from the trunk, with 400 000 axons entering the brain, the VsensR/IOM should be highly peripherin-immunopositive, including the myelinated axon bundles?! In this sense, the authors argue against their own interpretation - either the elephant trunk is not a highly sensitive tactile organ, or the VsensR is not the trigeminal nuclei it is supposed to be.

Comment: We made sure that elephant climbing fibers are strongly peripherin-positive (our revised Figure 2). As we noted in already our previous ms, we see weak diffuse peripherin-reactivity in the trigeminal nucleus (the inferior olive according to the referee), but no peripherin-reactive axon bundles (i.e. climbing fibers) that are seen in the inferior olive of other species. We also see no peripherin-reactive axon bundles (i.e. the olivo-cerebellar tract) arriving in the trigeminal nucleus as the tissue surrounding the trigeminal nucleus is devoid of peripherin-reactivity. Again, this finding is incompatible with the referee’s ideas. As far as we can tell, the trigeminal fibers are not reactive for peripherin in the elephant, i.e. we did not observe peripherin-reactivity very close to the nerve entry, but unfortunately, we did not stain for peripherin-reactivity into the nerve. As the referee alludes to the absence of peripherin-reactivity in the trigeminal tract is a difference between rodents and elephants.

Change: Our novel Figure 2.

Summary:

(1) Comparative data of species closely related to elephants (Afrotherians) demonstrates that not all mammals exhibit the "serrated" appearance of the principal nucleus of the inferior olive.

(2) The location of the IO and Vsens as reported in the current study (IOR and VsensR) would require a significant, and unprecedented, rearrangement of the brainstem in the elephants independently. I argue that the underlying molecular and genetic changes required to achieve this would be so extreme that it would lead to lethal phenotypes. Arguing that the "switcheroo" of the IO and Vsens does occur in the elephant (and no other mammals) and thus doesn't lead to lethal phenotypes is a circular argument that cannot be substantiated.

(3) Myelin stripes in the subnuclei of the inferior olivary nuclear complex are seen across all related mammals as shown above. Thus, the observation made in the elephant by the authors in what they call the VsensR, is similar to that seen in the IO of related mammals, especially when the IO takes on a more bulbous appearance. These myelin stripes are the origin of the olivocerebellar pathway, and are indeed calretinin immunopositive in the elephant as I show.

(4) What the authors see aligns perfectly with what has been described previously, the only difference being the names that nuclear complexes are being called. But identifying these nuclei is important, as any functional sequelae, as extensively discussed by the authors, is entirely dependent upon accurately identifying these nuclei.

(4) The peripherin immunostaining scores an own goal - if peripherin is marking peripheral nerves (as the authors and I believe it is), then why is the VsensR/IOM only "weakly positive" for this stain? This either means that the "extraordinary" tactile sensitivity of the elephant trunk is non-existent, or that the authors have misinterpreted this staining. That there is extensive staining in the fibre pathway dorsal and lateral to the IOR (which I call the spinal trigeminal tract), supports the idea that the authors have misinterpreted their peripherin immunostaining.

(5) Evolutionary expediency. The authors argue that what they report is an expedient way in which to modify the organisation of the brainstem in the elephant to accommodate the "extraordinary" tactile sensitivity. I disagree. As pointed out in my first review, the elephant cerebellum is very large and comprised of huge numbers of morphologically complex neurons. The inferior olivary nuclei in all mammals studied in detail to date, give rise to the climbing fibres that terminate on the Purkinje cells of the cerebellar cortex. It is more parsimonious to argue that, in alignment with the expansion of the elephant cerebellum (for motor control of the trunk), the inferior olivary nuclei (specifically the principal nucleus) have had additional neurons added to accommodate this cerebellar expansion. Such an addition of neurons to the principal nucleus of the inferior olive could readily lead to the loss of the serrated appearance of the principal nucleus of the inferior olive, and would require far less modifications in the developmental genetic program that forms these nuclei. This type of quantitative change appears to be the primary way in which structures are altered in the mammalian brainstem.

Comment: We still disagree with the referee. We note that our conclusions rest on the analysis of 8 elephant brainstems, which we sectioned in three planes and stained with a variety of metabolic and antibody stains and in which assigned two structures (the inferior olive and the trigeminal nucleus). Most of the evidence cited by the referee stems from a single paper, in which 147 structures were identified based on the analysis of a single brainstem sectioned in one plane and stained with a limited set of antibodies. Our synopsis of the evidence is the following.

(1) We agree with the referee that concerning brainstem position our scheme of a ventromedial trigeminal nucleus and a dorsolateral inferior olive deviates from the usual mammalian position of these nuclei (i.e. a dorsolateral trigeminal nucleus and a ventromedial inferior olive).

(2) Cytoarchitectonics support our partitioning scheme. The compact cellular appearance of our ventromedial trigeminal nucleus is characteristic of trigeminal nuclei. The serrated appearance of our dorsolateral inferior olive is characteristic of the mammalian inferior olive; we acknowledge that the referee claims exceptions here. To our knowledge, nobody has described a mammalian trigeminal nucleus with a serrated appearance (which would apply to the elephant in case the trigeminal nucleus is situated dorsolaterally).

(3) Metabolic staining (Cyto-chrome-oxidase reactivity) supports our partitioning scheme. Specifically, our ventromedial trigeminal nucleus shows intense Cyto-chrome-oxidase reactivity as it is seen in the trigeminal nuclei of trigeminal tactile experts.

(4) Isomorphism. The myelin stripes on our ventromedial trigeminal nucleus are isomorphic to trunk wrinkles. Isomorphism is a characteristic of somatosensory brain structures (barrel, barrelettes, nose-stripes, etc) and we know of no case, where such isomorphism was misleading.

(5) The large-scale organization of our ventromedial trigeminal nuclei in anterior-posterior repeats is characteristic of the mammalian trigeminal nuclei. To our knowledge, no such organization has ever been reported for the inferior olive.

(6) Connectivity analysis supports our partitioning scheme. According to our delineation of the elephant olivo-cerebellar tract, our dorsolateral inferior olive is connected via peripherin-positive climbing fibers to the cerebellum. In contrast, our ventromedial trigeminal nucleus (the referee’s inferior olive) is not connected via climbing fibers to the cerebellum.

Change: As discussed, we advanced further evidence in this revision. Our partitioning scheme (a ventromedial trigeminal nucleus and a dorsolateral inferior olive) is better supported by data and makes more sense than the referee’s suggestion (a dorsolateral trigeminal nucleus and a ventromedial inferior olive). It should be published.

Reviewer #3 (Public Review):

Summary:

The study claims to investigate trunk representations in elephant trigeminal nuclei located in the brainstem. The researchers identify large protrusions visible from the ventral surface of the brainstem, which they examined using a range of histological methods. However, this ventral location is usually where the inferior olivary complex is found, which challenges the author's assertions about the nucleus under analysis. They find that this brainstem nucleus of elephants contains repeating modules, with a focus on the anterior and largest unit which they define as the putative nucleus principalis trunk module of the trigeminal. The nucleus exhibits low neuron density, with glia outnumbering neurons significantly. The study also utilizes synchrotron X-ray phase contrast tomography to suggest that myelin-stripe-axons traverse this module. The analysis maps myelin-rich stripes in several specimens and concludes that based on their number and patterning that they likely correspond with trunk folds; however this conclusion is not well supported if the nucleus has been misidentified.

Comment: The referee provides a summary of our work. The referee also notes that the correct identification of the trigeminal nucleus is critical to the message of our paper.

Change: In line with these assessments we focused our revision efforts on the issue of trigeminal nucleus identification, please see our introductory comments and our response to Referee 2.

Strengths:

The strength of this research lies in its comprehensive use of various anatomical methods, including Nissl staining, myelin staining, Golgi staining, cytochrome oxidase labeling, and synchrotron X-ray phase contrast tomography. The inclusion of quantitative data on cell numbers and sizes, dendritic orientation and morphology, and blood vessel density across the nucleus adds a quantitative dimension. Furthermore, the research is commendable for its high-quality and abundant images and figures, effectively illustrating the anatomy under investigation.

Comment: We appreciate this positive assessment.

Change: None

Weaknesses:

While the research provides potentially valuable insights if revised to focus on the structure that appears to be inferior olivary nucleus, there are certain additional weaknesses that warrant further consideration. First, the suggestion that myelin stripes solely serve to separate sensory or motor modules rather than functioning as an "axonal supply system" lacks substantial support due to the absence of information about the neuronal origins and the termination targets of the axons. Postmortem fixed brain tissue limits the ability to trace full axon projections. While the study acknowledges these limitations, it is important to exercise caution in drawing conclusions about the precise role of myelin stripes without a more comprehensive understanding of their neural connections.

Comment: We understand these criticisms and the need for cautious interpretation. As we noted previously, we think that the Elife-publishing scheme, where critical referee commentary is published along with our ms, will make this contribution particularly valuable.

Change: Our additional efforts to secure the correct identification of the trigeminal nucleus.

Second, the quantification presented in the study lacks comparison to other species or other relevant variables within the elephant specimens (i.e., whole brain or brainstem volume). The absence of comparative data to different species limits the ability to fully evaluate the significance of the findings. Comparative analyses could provide a broader context for understanding whether the observed features are unique to elephants or more common across species. This limitation in comparative data hinders a more comprehensive assessment of the implications of the research within the broader field of neuroanatomy. Furthermore, the quantitative comparisons between African and Asian elephant specimens should include some measure of overall brain size as a covariate in the analyses. Addressing these weaknesses would enable a richer interpretation of the study's findings.

Comment: We understand, why the referee asks for additional comparative data, which would make our study more meaningful. We note that we already published a quantitative comparison of African and Asian elephant facial nuclei (Kaufmann et al. 2022). The quantitative differences between African and Asian elephant facial nuclei are similar in magnitude to what we observed here for the trigeminal nucleus, i.e. African elephants have about 10-15% more facial nucleus neurons than Asian elephants. The referee also notes that data on overall elephant brain size might be important for interpreting our data. We agree with this sentiment and we are preparing a ms on African and Asian elephant brain size. We find – unexpectedly given the larger body size of African elephants – that African elephants have smaller brains than Asian elephants. The finding might imply that African elephants, which have more facial nucleus neurons and more trigeminal nucleus trunk module neurons, are neurally more specialized in trunk control than Asian elephants.

Change: We are preparing a further ms on African and Asian elephant brain size, a first version of this work has been submitted.

Reviewer #4 (Public Review):

Summary:

The authors report a novel isomorphism in which the folds of the elephant trunk are recognizably mapped onto the principal sensory trigeminal nucleus in the brainstem. Further, they identifiy the enlarged nucleus as being situated in this species in an unusual ventral midline position.

Comment: The referee summarizes our work.

Change: None.

Strengths:

The identity of the purported trigeminal nucleus and the isomorphic mapping with the trunk folds is supported by multiple lines of evidence: enhanced staining for cytochrome oxidase, an enzyme associated with high metabolic activity; dense vascularization, consistent with high metabolic activity; prominent myelinated bundles that partition the nucleus in a 1:1 mapping of the cutaneous folds in the trunk periphery; near absence of labeling for the anti-peripherin antibody, specific for climbing fibers, which can be seen as expected in the inferior olive; and a high density of glia.

Comment: The referee again reviews some of our key findings.

Change: None.

Weaknesses:

Despite the supporting evidence listed above, the identification of the gross anatomical bumps, conspicuous in the ventral midline, is problematic. This would be the standard location of the inferior olive, with the principal trigeminal nucleus occupying a more dorsal position. This presents an apparent contradiction which at a minimum needs further discussion. Major species-specific specializations and positional shifts are well-documented for cortical areas, but nuclear layouts in the brainstem have been considered as less malleable.

Comment: The referee notes that our discrepancy with referee 2, needs to be addressed with further evidence and discussion, given the unusual position of both inferior olive and trigeminal nucleus in the partitioning scheme and that the mammalian brainstem tends to be positionally conservative. We agree with the referee. We note that – based on the immense size of the elephant trigeminal ganglion (50 g), half the size of a monkey brain – it was expected that the elephant trigeminal nucleus ought to be exceptionally large.

Change: We did additional experimental work to resolve this matter: (i) We ascertained that elephant climbing fibers are strongly peripherin-positive. (ii) Based on elephant climbing fiber peripherin-reactivity we delineated the elephant olivo-cerebellar tract. We find that the olivo-cerebellar connects to the structure we refer to as inferior olive to the cerebellum. (iii) We also found that the trigeminal nucleus (the structure the referee refers to as inferior olive) appears to receive no climbing fibers. (iv) We provide indications that the tracing of the trigeminal nerve into the olivo-cerebellar tract by Maseko et al. 2023 was erroneous (Referee-Figure 1). These novel findings support our ideas.

Reviewer #5 (Public Review):

After reading the manuscript and the concerns raised by reviewer 2 I see both sides of the argument - the relative location of trigeminal nucleus versus the inferior olive is quite different in elephants (and different from previous studies in elephants), but when there is a large disproportionate magnification of a behaviorally relevant body part at most levels of the nervous system (certainly in the cortex and thalamus), you can get major shifting in location of different structures. In the case of the elephant, it looks like there may be a lot of shifting. Something that is compelling is that the number of modules separated but the myelin bands correspond to the number of trunk folds which is different in the different elephants. This sort of modular division based on body parts is a general principle of mammalian brain organization (demonstrated beautifully for the cuneate and gracile nucleus in primates, VP in most of species, S1 in a variety of mammals such as the star nosed mole and duck-billed platypus). I don't think these relative changes in the brainstem would require major genetic programming - although some surely exists. Rodents and elephants have been independently evolving for over 60 million years so there is a substantial amount of time for changes in each l lineage to occur.

I agree that the authors have identified the trigeminal nucleus correctly, although comparisons with more out groups would be needed to confirm this (although I'm not suggesting that the authors do this). I also think the new figure (which shows previous divisions of the brainstem versus their own) allows the reader to consider these issues for themselves. When reviewing this paper, I actually took the time to go through atlases of other species and even look at some of my own data from highly derived species. Establishing homology across groups based only on relative location is tough especially when there appears to be large shifts in relative location of structures. My thoughts are that the authors did an extraordinary amount of work on obtaining, processing and analyzing this extremely valuable tissue. They document their work with images of the tissue and their arguments for their divisions are solid. I feel that they have earned the right to speculate - with qualifications - which they provide.

Comment: The referee summarizes our work and appears to be convinced by the line of our arguments. We are most grateful for this assessment. We add, again, that the skeptical assessment of referee 2 will be published as well and will give the interested reader the possibility to view another perspective on our work.

Change: None.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

With this manuscript being virtually identical to the previous version, it is possible that some of the definitive conclusions about having identified the elephant trigeminal nucleus and trunk representation should be moderated in a more nuanced manner, especially given the careful and experienced perspective from reviewers with first hand knowledge elephant neuroanatomy.

Comment: We agree that both our first and second revisions were very much centered on the debate of the correct identification of the trigeminal nucleus and that our ms did not evolve as much in other regards. This being said we agree with Referee 2 that we needed to have this debate. We also think we advanced important novel data in this context (the delineation of elephant olivo-cerebellar tract through the peripherin-antibody).

Changes: Our revised Figure 2.

The peripherin staining adds another level of argument to the authors having identified the trigeminal brainstem instead of the inferior olive, if differential expression of peripherin is strong enough to distinguish one structure from the other.

Comment: We think we showed too little peripherin-antibody staining in our previous revision. We have now addressed this problem.

Changes: Our revised Figure 2, i.e. the delineation of elephant olivo-cerebellar tract through the peripherin-antibody).

There are some minor corrections to be made with the addition of Fig. 2., including renumbering the figures in the manuscript (e.g., 406, 521).

I continue to appreciate this novel investigation of the elephant brainstem and find it an interesting and thorough study, with the use of classical and modern neuroanatomical methods.

Comment: We are thankful for this positive assessment.

Reviewer #2 (Recommendations For The Authors):

I do realise the authors are very unhappy with me and the reviews I have submitted. I do apologise if feelings have been hurt, and I do understand the authors put in a lot of hard work and thought to develop what they have; however, it is unfortunate that the work and thoughts are not correct. Science is about the search for the truth and sometimes we get it wrong. This is part of the scientific process and why most journals adhere to strict review processes of scientific manuscripts. As I said previously, the authors can use their data to write a paper describing and quantifying Golgi staining of neurons in the principal olivary nucleus of the elephant that should be published in a specialised journal and contextualised in terms of the motor control of the trunk and the large cerebellum of the elephant.

Comment: We appreciate the referee’s kind words. Also, no hard feelings from our side, this is just a scientific debate. In our experience, neuroanatomical debates are resolved by evidence and we note that we provide evidence strengthening our identification of the trigeminal nucleus and inferior olive. As far as we can tell from this effort and the substantial evidence accumulated, the referee is wrong.

Reviewer #4 (Recommendations For The Authors):

As a new reviewer, I have benefited from reading the previous reviews and Author response, even while having several new comments to add.

(1) The identification of the inferior olive and trigeminal nuclei is obviously center stage. An enlargement of the trigeminal nuclei is not necessarily problematic, given the published reports on the dramatic enlargement of the trigeminal nerve (Purkart et al., 2022). At issue is the conspicuous relocation of the trigeminal nuclei that is being promoted by Reveyaz et al. Conspicuous rearrangements are not uncommon; for example, primary sensory cortical fields in different species (fig. 1 in H.H.A. Oelschlager for dolphins; S. De Vreese et al. (2023) for cetaceans, L. Krubitzer on various species, in the context of evolution). The difficult point here concerns what looks like a rather conspicuous gross anatomical rearrangement, in BRAINSTEM - the assumption being that the brainstem bauplan is going to be specifically conservative and refractory to gross anatomical rearrangement.

Comment: We agree with the referee that the brainstem rearrangements are unexpected. We also think that the correct identification of nuclei needs to be at the center of our revision efforts.

Change: Our revision provided further evidence (delineation of the olivo-cerebellar tract, characterization of the trigeminal nerve entry) about the identity of the nuclei we studied.

Why would a major nucleus shift to such a different location? and how? Can ex vivo DTI provide further support of the correct identification? Is there other "disruption" in the brainstem? What occupies the traditional position of the trigeminal nuclei? An atlas-equivalent coronal view of the entire brainstem would be informative. The Authors have assembled multiple criteria to support their argument that the ventral "bumps" are in fact a translocated trigeminal principal nucleus: enhanced CO staining, enhanced vascularization, enhanced myelination (via Golgi stains and tomography), very scant labeling for a climbing fiber specific antibody ( anti-peripherin), vs. dense staining of this in the alternative structure that they identify as IO; and a high density of glia. Admittedly, this should be sufficient, but the proposed translocation (in the BRAINSTEM) is sufficiently startling that this is arguably NOT sufficient.

The terminology of "putative" is helpful, but a more cogent presentation of the results and more careful discussion might succeed in winning over at least some of a skeptical readership.

Comment: We do not know, what led to the elephant brainstem rearrangements we propose. If the trigeminal nuclei had expanded isometrically in elephants from the ancestral pattern, one would have expected a brain with big lateral bumps, not the elephant brain with its big ventromedial bumps. We note, however, that very likely the expansion of the elephant trigeminal nuclei did not occur isometrically. Instead, the neural representation of the elephant nose expanded dramatically and in rodents the nose is represented ventromedially in the brainstem face representation. Thus, we propose a ‘ventromedial outgrowth model’ according to which the elephant ventromedial trigeminal bumps result from a ventromedially direct outgrowth of the ancestral ventromedial nose representation.

We advanced substantially more evidence to support our partitioning scheme, including the delineation of the olivo-cerebellar tract based on peripherin-reactivity. We also identified problems in previous partitioning schemes, such as the claim that the trigeminal nerve continues into the ~4x smaller olivocerebellar tract (Referee-Figure 1C, D); we think such a flow of fibers, (which is also at odds with peripherin-antibody-reactivity and the appearance of nerve and olivocerebellar tract), is highly unlikely if not physically impossible. With all that we do not think that we overstate our case in our cautiously presented ms.

Change: We added evidence on the identification of elephant trigeminal nuclei and inferior olive.

(2) Role of myelin. While the photos of myelin are convincing, it would be nice to have further documentation. Gallyas? Would antibodies to MBP work? What is the myelin distribution in the "standard" trigeminal nuclei (human? macaque or chimpanzee?). What are alternative sources of the bundles? Regardless, I think it would be beneficial to de-emphasize this point about the role of myelin in demarcating compartments.

I would in fact suggest an alternative (more neutral) title that might highlight instead the isomorphic feature; for example, "An isomorphic representation of Trunk folds in the Elephant Trigeminal Nucleus." The present title stresses myelin, but figure 1 already focuses on CO. Additionally, the folds are actually mentioned almost in passing until later in the manuscript. I recommend a short section on these at the beginning of the Results to serve as a useful framework.

Here I'm inclined to agree with the Reviewer, that the Authors' contention that the myelin stipes serve PRIMARILY to separate trunk-fold domains is not particularly compelling and arguably a distraction. The point can be made, but perhaps with less emphasis. After all, the fact that myelin has multiple roles is well-established, even if frequently overlooked. In addition, the Authors might make better use of an extensive relevant literature related to myelin as a compartmental marker; for example, results and discussion in D. Haenelt....N. Weiskopf (eLife, 2023), among others. Another example is the heavily myelinated stria of Gennari in primate visual cortex, consisting of intrinsic pyramidal cell axons, but where the role of the myelination has still not been elucidated.

Comment: (1) Documentation of myelin. We note that we show further identification of myelinated fibers by the fluorescent dye fluomyelin in Figure 4B. We also performed additional myelin stains as the gold-myelin stain after the protocol of Schmued (Referee-Figure 2). In the end, nothing worked quite as well to visualize myelin-stripes as the bright-field images shown in Figure 4A and it is only the images that allowed us to match myelin-stripes to trunk folds. Hence, we focus our presentation on these images.

(2) Title: We get why the referee envisions an alternative title. This being said, we would like to stick with our current title, because we feel it highlights the major novelty we discovered.

(3) We agree with many of the other comments of the referee on myelin phenomenology. We missed the Haenelt reference pointed out by the referee and think it is highly relevant to our paper

Change: 1. Review image 2. Inclusion of the Haenelt-reference.

Myelin stripes of the elephant trunk module visualized by Gold-chloride staining according to Schmued.

A, Low magnification micrograph of the trunk module of African elephant Indra stained with AuCl according to Schmued. The putative finger is to the left, proximal is to the right. Myelin stripes can easily be recognized. The white box indicates the area shown in B. B, high magnification micrograph of two myelin stripes. Individual gold-stained (black) axons organized in myelin stripes can be recognized.

Schmued, L. C. (1990). A rapid, sensitive histochemical stain for myelin in frozen brain sections. Journal of Histochemistry & Cytochemistry,38(5), 717-720.

Are the "bumps" in any way "analogous" to the "brain warts" seen in entorhinal areas of some human brains (G. W. van Hoesen and A. Solodkin (1993)?

Comment: We think this is a similar phenomenon.

Change: We included the Hoesen and A. Solodkin (1993) reference in our discussion.

At least slightly more background (ie, a separate section or, if necessary, supplement) would be helpful, going into more detail on the several subdivisions of the ION and if these undergo major alterations in the elephant.

Comment: The strength of the paper is the detailed delineation of the trunk module, based on myelin stripes and isomorphism. We don’t think we have strong evidence on ION subdivisions, because it appears the trigeminal tract cannot be easily traced in elephants. Accordingly, we find it difficult to add information here.

Change: None.

Is there evidence from the literature of other conspicuous gross anatomical translocations, in any species, especially in subcortical regions?

Comment: The best example that comes to mind is the star-nosed mole brainstem. There is a beautiful paper comparing the star-nosed mole brainstem to the normal mole brainstem (Catania et al 2011). The principal trigeminal nucleus in the star-nosed mole is far more rostral and also more medial than in the mole; still, such rearrangements are minor compared to what we propose in elephants.

Catania, Kenneth C., Duncan B. Leitch, and Danielle Gauthier. "A star in the brainstem reveals the first step of cortical magnification." PloS one 6.7 (2011): e22406.

Change: None.

(3) A major point concerns the isomorphism between the putative trigeminal nuclei and the trunk specialization. I think this can be much better presented, at least with more discussion and other examples. The Authors mention about the rodent "barrels," but it seemed strange to me that they do not refer to their own results in pig (C. Ritter et al., 2023) nor the work from Ken Catania, 2002 (star-nosed mole; "fingerprints in the brain") or other that might be appropriate. I concur with the Reviewer that there should be more comparative data.

Comment: We agree.

Change: We added a discussion of other isomorphisms including the the star-nosed mole to our paper.

(4) Textual organization could be improved.

The Abstract all-important Introduction is a longish, semi "run-on" paragraph. At a minimum this should be broken up. The last paragraph of the Introduction puts forth five issues, but these are only loosely followed in the Results section. I think clarity and good organization is of the upmost importance in this manuscript. I recommend that the Authors begin the Results with a section on the trunk folds (currently figure 5, and discussion), continue with the several points related to the identification of the trigeminal nuclei, and continue with a parallel description of ION with more parallel data on the putative trigeminal and IO structures (currently referee Table 1, but incorporate into the text and add higher magnification of nucleus-specific cell types in the IO and trigeminal nuclei). Relevant comparative data should be included in the Discussion.

Comment: 1. We agree with the referee that our abstract needed to be revised. 2. We also think that our ms was heavily altered by the insertion of the new Figure 2, which complemented Figure 1 from our first submission and is concerned with the identification of the inferior olive. From a standpoint of textual flow such changes were not ideal, but the revisions massively added to the certainty with which we identify the trigeminal nuclei. Thus, although we are not as content as we were with the flow, we think the ms advanced in the revision process and we would like to keep the Figure sequence as is. 3. We already noted above that we included additional comparative evidence.

Change: 1. We revised our abstract. 2. We added comparative evidence.

Reviewer #5 (Recommendations For The Authors):

The data is invaluable and provides insights into some of the largest mammals on the planet.

Comment: We are incredibly thankful for this positive assessment.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Elephants have elaborate trunk skills and large, but poorly understood brains. Here we study trunk representations in elephant trigeminal nuclei, which form large protrusions on the ventral brainstem. Dense vascularization and intense cytochrome-oxidase reactivity distinguish several elongated putative trunk modules, which repeat in the anterior-posterior direction; our analysis focuses on the most anterior and largest of the units, the putative nucleus principalis trunk module. Module neuron density is low and non-neural cells outnumber neurons by ∼108:1. Dendritic trees are elongated along the axis of axon bundles (myelin stripes) transversing the trunk module. Furthermore, synchrotron X-ray phase contrast tomography suggests myelin-stripe-axons transverse the trunk module. We show a remarkable correspondence of trunk module myelin stripes and trunk folds. Myelin stripes show little relation to trigeminal neurons and stripe-axons appear to often go nowhere; these observations suggest the possibility that myelin stripes might serve to separate trunk-fold domains rather than to connect neurons. The myelin-stripes-to-folds mapping allowed to determine neural magnification factors, which changed from 1000:1 proximally to 5:1 in the trunk finger. Asian elephants have fewer (∼640,000) trunk-module neurons than Africans (∼740,000) and show enlarged representations of trunk parts involved in object wrapping. We conclude the elephant trigeminal trunk module is exquisitely organized into trunk-fold-related units.

Revised version: Elephants have elaborate trunk skills and large, but poorly understood brains. Here we study trunk representations in elephant trigeminal nuclei, which form large protrusions on the ventral brainstem. These ventral brainstem protrusions have previously been referred to as inferior olive, but a delineation of the olivo-cerebellar tract reveals these (trigeminal) nuclei are not connected to the cerebellum via climbing fibers. In contrast, the olivo-cerebellar tract connects to a large dorsolateral nucleus with a serrated cellular architecture, the putative elephant inferior olive. Dense vascularization and intense cytochrome-oxidase reactivity distinguish several elongated trigeminal putative trunk modules, which repeat in the anterior-posterior direction. We focus on the most anterior and largest of these units, the putative nucleus principalis trunk module. Module neuron density is low and non-neural cells outnumber neurons by ∼108:1. Dendritic trees are elongated along the axis of axon bundles (myelin stripes) transversing the trunk module. Synchrotron X-ray-phase-contrast tomography suggests myelin-stripe-axons transverse the trunk module. We show a remarkable correspondence of trunk module myelin stripes and trunk folds. Myelin stripes show little relation to trigeminal neurons and stripe-axons appear to often go nowhere; we suggest that myelin stripes might serve to separate trunk-fold domains rather than to connect neurons. Myelin-stripes-to-folds mapping allowed to determine neural magnification factors, which changed from 1000:1 proximally to 5:1 in the trunk finger. Asian elephants have fewer (∼640,000) trunk-module neurons than Africans (∼740,000) and show enlarged representations of trunk parts involved in object wrapping. The elephant trigeminal trunk module is exquisitely organized into trunk-fold-related units.

Introduction

Earlier version: Elephants are the largest extant terrestrial animals and rely on their trunks to acquire huge amounts of food. The trunk is a fusion organ of the nose and upper lip. Nose lip fusion occurs in the fourth month of elephant fetal development (Fischer & Trautmann, 1987; Schulz et al., 2023). The trunk is an immensely muscular structure (Cuvier, 1796; Cuvier and Dumeril, 1838; Shoshani, 1982) that contains about 90,000 muscle fascicles (Longren et al., 2023). Not surprisingly, the trunk’s prime motor control structure, the facial nucleus, is very large (Maseko et al., 2013) and shows an elaborate cellular architecture (Kaufmann et al., 2022). Trunks have prominent folds that differ between African (Loxodonta africana) and Asian (Elephas maximus) elephants (Schulz et al., in prep). Interestingly, the object-grasping behavior of African and Asian elephants differs considerably. Asian elephants have a single trunk finger and tend to wrap their trunk around objects, whereas African elephants prefer pinching objects with their two trunk fingers (Racine, 1980). Elephants are skillful with their trunks (Shoshani, 1992), have a high tactile sensitivity (Dehnhardt, Friese, and Sachser, 1997), and even acquire dexterous manipulative behaviors such as banana peeling (Kaufmann et al., 2023). Tactile feedback is of great significance for trunk behaviors because elephants have only limited visual abilities. Recently, Deiringer et al. (2023) investigated trunk whiskers and observed use-dependent whisker lateralization, dense whisker arrays on the trunk tip and the ventral trunk, and marked whisker differences between African and Asian elephants. The sensory periphery of elephant trunks was investigated in a landmark study by Rasmussen & Munger (1996), who described dense innervation patterns in the elephant fingertip. The whole elephant trunk is massively innervated by large trigeminal ganglia (Sprinz, 1952; Purkart et al., 2022). The cerebellum of elephants is very large both in relative and absolute terms (Maseko et al., 2012) and it is believed to be an important part of the control of the trunk movement, as a sensory and motor processing area (Maseko et al., (2013). The turning point in the investigation of the mammalian trigeminal system was the description of the cortical whisker barrels by Woolsey & Van der Loss (1970 and this work informed our approach to the elephant brainstem. The recognition of the cortical barrel pattern has led to thousands of follow-up studies, which included the discovery of whisker-related thalamic so-called barreloids (Van der Loos, 1976) and whisker-related units in the trigeminal brainstem, so-called barrelettes (Belford & Killackey, 1979; Ma, 1991). As shown by Belford & Killackey (1979) and Ma (1991) the brainstem contains several topographic trigeminal representations, which repeat in anterior to posterior direction. Specifically, these studies identified the most anterior one as the largest sensory trigeminal representation (the nucleus principalis) and several smaller more posterior trigeminal sensory nuclei. We will adopt the trigeminal terminology established by these authors (Belford & Killackey, 1979; Ma, 199).

Revised version: Elephants are the largest extant terrestrial animals and rely on their trunks to acquire huge amounts of food. The trunk is a fusion organ of the nose and upper lip. Nose lip fusion occurs in the fourth month of elephant fetal development (Fischer & Trautmann, 1987; Schulz et al., 2023). The trunk is an immensely muscular structure (Cuvier, 1796; Cuvier and Dumeril, 1838; Shoshani, 1982) that contains about 90,000 muscle fascicles (Longren et al., 2023). Not surprisingly, the trunk’s prime motor control structure, the facial nucleus, is very large (Maseko et al., 2013) and shows an elaborate cellular architecture (Kaufmann et al., 2022). Trunks have prominent folds that differ between African (Loxodonta africana) and Asian (Elephas maximus) elephants (Schulz et al., in prep). Interestingly, the object-grasping behavior of African and Asian elephants differs considerably. Asian elephants have a single trunk finger and tend to wrap their trunk around objects, whereas African elephants prefer pinching objects with their two trunk fingers (Racine, 1980). Elephants are skillful with their trunks (Shoshani, 1992), have a high tactile sensitivity (Dehnhardt, Friese, and Sachser, 1997), and even acquire dexterous manipulative behaviors such as banana peeling (Kaufmann et al., 2023). Tactile feedback is of great significance for trunk behaviors because elephants have only limited visual abilities. Recently, Deiringer et al. (2023) investigated trunk whiskers and observed use-dependent whisker lateralization, dense whisker arrays on the trunk tip and the ventral trunk, and marked whisker differences between African and Asian elephants. The sensory periphery of elephant trunks was investigated in a landmark study by Rasmussen & Munger (1996), who described dense innervation patterns in the elephant fingertip. The whole elephant trunk is massively innervated by large trigeminal ganglia (Sprinz, 1952; Purkart et al., 2022). The cerebellum of elephants is very large both in relative and absolute terms (Maseko et al., 2012) and it is believed to be an important part of the control of the trunk movement, as a sensory and motor processing area (Maseko et al., (2013). The turning point in the investigation of the mammalian trigeminal system was the description of the cortical whisker barrels by Woolsey & Van der Loss (1970 and this work informed our approach to the elephant brainstem. The recognition of the cortical barrel pattern has led to thousands of follow-up studies, which included the discovery of whisker-related thalamic so-called barreloids (Van der Loos, 1976) and whisker-related units in the trigeminal brainstem, so-called barrelettes (Belford & Killackey, 1979; Ma, 1991). As shown by Belford & Killackey (1979) and Ma (1991) the brainstem contains several topographic trigeminal representations, which repeat in anterior to posterior direction. Specifically, these studies identified the most anterior one as the largest sensory trigeminal representation (the nucleus principalis) and several smaller more posterior trigeminal sensory nuclei. We will adopt the trigeminal terminology established by these authors (Belford & Killackey, 1979; Ma, 1991).

A variety of excellent studies have investigated the cellular statistics of elephant brains and indicated elephant brains are not simply scaled-up mouse brains. Specifically, investigators found much lower neuronal densities in elephant brains than in rodents (Haug, 1987; Herculano-Houzel et al., 2014). A prominent difference between small and large brains is the increased amounts of white matter in larger brains. Myelin sheaths, which give white matter its whitish shine, and the enwrapped axons are usually thought to form a supply and connectivity system, an idea we will question in our study.

We pursued the following questions: (i) Can candidate regions for the elephant trigeminal trunk representation be identified? (ii) If yes, can multiple sensory trigeminal nuclei be identified as in other mammals? (iii) What is the neuroanatomical structure of the elephant brainstem trunk representations? (iv) Do the elaborate myelin structures in the elephant trigeminal nuclei form an axonal supply system? (v) How does the organization of elephant brainstem trunk representations relate to the differential trunk morphology and grasping behavior of African and Asian elephants? We tentatively identified an elephant brainstem trunk module characterized by intense metabolism and vascularization. The putative trunk module contains a myelin map of trunk folds. This myelin map allows precise mapping of the neural topography of the trunk representation and reveals species differences between African and Asian elephants.

Results

Overview

Earlier version: Determining trigeminal representations in elephants is challenging because invasive recordings or invasive viral tracing methods cannot be applied. We proceeded to build a hypothesis on the elephant trigeminal brainstem trunk in six steps. First, we identified a candidate module for the brainstem trunk representation. Second, we showed that this module architecture repeats in the anterior-posterior direction in the elephant brainstem. Third, we characterized the cellular organization of this putative trunk module. Fourth, we documented a close correspondence between the myeloarchitecture of this module and the folds of the elephant trunk. Fifth, we applied synchrotron X-ray tomography to assess the microscopic architecture of myelin stripes. Sixth, we showed that species-specific differences in trunk structure have correlates in the putative trunk module.

Revised version: Determining trigeminal representations in elephants is challenging because invasive recordings or invasive viral tracing methods cannot be applied. We proceeded to build a hypothesis on the elephant trigeminal brainstem trunk in seven steps. First, we identified a candidate module for the brainstem trunk representation. Second, we used peripherin-antibody staining to delineate the elephant olivo-cerebellar tract. This analysis indicated to us the ventral brainstem bumps do not correspond to the elephant inferior olive as previously thought (Maseko et al. 2013). Third, we showed that this module architecture repeats in the anterior-posterior direction in the elephant brainstem. Fourth, we characterized the cellular organization of this putative trunk module. Fifth, we documented a close correspondence between the myeloarchitecture of this module and the folds of the elephant trunk. Sixth, we applied synchrotron X-ray tomography to assess the microscopic architecture of myelin stripes. Seventh, we showed that species-specific differences in trunk structure have correlates in the putative trunk module.

A metabolically highly active, strongly vascularized putative trunk module

The brain of the Asian elephant cow Burma is shown in Figure 1A. In rodents, the sensory trigeminal nuclei are observed posterior to the pons. As shown in Figure 1B, in a ventral view of the brain stem of Burma, about 1 cm posterior to the (large) pons, a pair of large bumps is obvious on the ventral brainstem surface. By size and pronounced protrusion, these bumps on the ventral brainstem distinguish elephant brains from those of other mammals. Much smaller bumps are seen in a similar position in the human brain, where they contain the inferior olive. Accordingly, the bumps of the elephant brain have been referred to as the elephant inferior olive (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962), but our investigation did not support this idea. We sectioned the elephant medulla and stained sections for cytochrome oxidase reactivity, a mitochondrial enzyme, the activity of which is closely related to tissue energy consumption. Trigeminal nuclei tend to show intense activity in cytochrome oxidase reactivity (Belford & Killackey, 1979; Ma, 1991). A cytochrome oxidase-stained coronal section through the bump of the Asian elephant bull Raj is shown in Figure 1C. We found that the bump contained the most intense cytochrome oxidase reactivity in the elephant brainstem and (to the extent that we performed such staining in other brain regions) the rest of the elephant brain. Three cytochrome-reactive modules (a putative trunk module, a putative nostril module, and a putative lower lip/jaw) are obvious, the largest of which we refer to as a putative trunk module (Figure 1D). The putative trunk module is elongated and we hypothesize that a particularly intensely cytochrome oxidase reactive region at the ventrolateral pole of the module corresponds to the dorsal finger representation (Figure 1C, D). We provide a detailed justification for our assignments of a putative trunk module, a putative nostril module, and a putative lower lip/jaw trigeminal module in Figure 2. We also studied brain stem sections in the African elephant (Figure 1E) and identified a similar putative trunk module there. We investigated the vascularization of this module, which was evident from the cytochrome oxidase reactivity of erythrocytes in blood vessels of our non-perfused elephant brains (Figure 1F), and found that the trunk module stands out from the rest of the brainstem (Figure 1G). Specifically, it contains about twice as many blood vessels per volume as the remainder of the brainstem (Figure 1H). In parasagittal sections, the putative trunk module had a compact appearance much like the trigeminal nuclei of other mammals (Figure 1I). In parasagittal sections lateral to the putative trunk module we observed a nucleus with a very distinct banded cellular appearance (Figure 1J), a cellular architecture characteristic of the inferior olive of other mammals (Brodal et al. 1980). We conclude that the elephant brainstem contains a large, highly vascularized, and highly cytochrome-oxidase reactive elongated putative trunk module.

Appearance, metabolism, blood supply, and identification of the elephant trigeminal nuclei

A, ventral view of the brain of Burma, a 52-year-old female Asian elephant.

B, ventral view of the brainstem of Burma. Note the large pons (upper part of the photograph). Posterior to the pons a pair of prominent bumps are visible, the putative trigeminal nuclei (TN) of the elephant. The trigeminal nuclei bumps are observed, wherein in humans and other mammals the inferior olive is found and was referred to as inferior olive by previous authors (Shoshani et al., 2006; Maseko et al., 2013; Rasenberger, 2019). The inferior olive can be identified at an unusual lateral position in elephants, however (IO).

C, 60 µm coronal section through the trigeminal nuclei of Raj, a four-year-old elephant bull, stained for cytochrome-oxidase reactivity, a mitochondrial enzyme, the reactivity of which reflects constitutive metabolic activity. The trigeminal nuclei show some of the strongest cytochrome-oxidase reactivity (indicated by the brown color) in the elephant brain and strong cytochrome-oxidase reactivity is typical for the trigeminal nuclei of many mammals.

D, a drawing of putative trigeminal subnuclei stained in C. Note the compact shape of the trunk module, which is unlike the inferior olive of mammals.

E, drawing of a coronal section stained for cytochrome-oxidase reactivity through the brainstem of Indra a female African elephant, borders of nuclei are outlined.

F, upper, micrograph of cytochrome-oxidase reactivity in the putative trunk fingertip representation. Lower, drawing of cytochrome-oxidase reactive erythrocytes in blood vessels. Note the high density of vessels inside of the fingertip but not in the surrounding tissue.

G, drawing of the entire brainstem section. The putative trunk module stands out from the rest of the brainstem in terms of vessel density.

H, quantification of blood vessel length in various parts of the brainstem. Note that not all vessels contain erythrocytes and are stained, i.e., measures of blood vessel length are lower bound estimates.

I, Brightfield micrograph of a parasaggital section through the trigeminal trunk module. Note the compact cellular architecture.

J, Brightfield micrograph of a parasaggital section through the inferior olive. Note the banded cellular architecture that is characteristic of the mammalian inferior olive.

L = lateral; V = ventral; P = posterior.

Earlier version: The putative inferior olive but not the putative trigeminal nucleus contains peripherin-positive axon bundles (presumptive climbing fibers).

Revised version: Peripherin-antibodies reveal climbing fibers and the olivo-cerebellar tract connecting the cerebellum to the inferior olive but not to the trigeminal nucleus

Added in the revised version: A, ventral view of the elephant brainstem with key structures schematically superimposed.

Added in the revised version: B, schematic of the elephant brainstem. We indicate the positions of the peripherin-stained sections shown in C-F and our delineation of the olivo-cerebellar tract based on peripherin-reactive putative climbing fibers.

Earlier version: A, Overview picture of a brainstem section stained with anti-peripherin-antibodies (white color). Anti-peripherin-antibodies stain climbing fibers in a wide variety of mammals. The section comes from the posterior brainstem of African elephant cow Bibi; in this posterior region, both putative inferior olive and trigeminal nucleus are visible. Note the bright staining of the dorsolateral nucleus, the putative inferior olive according to Reveyaz et al., and the trigeminal nucleus according to Maseko et al., 2013.

Revised version: C, left, overview picture of a brainstem section stained with anti-peripherin-antibodies (white color). Anti-peripherin-antibodies stain climbing fibers in a wide variety of mammals. The section comes from the posterior brainstem of African elephant cow Bibi; in this posterior region, both putative inferior olive and trigeminal nucleus are visible. Note the bright staining of the dorsolateral nucleus, the putative inferior olive according to Reveyaz et al. (the trigeminal nucleus according to Maseko et al., 2013). C, upper right, high magnification view of the dorsolateral nucleus (corresponding to the upper red rectangle in A). Anti-peripherin-positive axon bundles (red stars, putative climbing fibers) are seen in support of the inferior olive hypothesis of Reveyaz et al. C, lower right, high magnification view of the ventromedial nucleus (corresponding to the lower red rectangle in A). The ventromedial nucleus is weakly positive for peripherin but contains no anti-peripherin-positive axon bundles (i.e. no putative climbing fibers). Note that myelin stripes – green stars, weakly visible as dark omissions – are clearly anti-peripherin-negative. The region around the ventromedial nucleus is devoid of peripherin-reactivity and this is true throughout the brainstem. This observation suggests that the putative trigeminal nucleus does not receive climbing fiber input.

Added in the revised version: D, anti-peripherin-antibody (bright color) stained section below the cerebellar peduncle. Putative climbing fibers (arrow) can be seen to be budding off the olivo-cerebellar (red star) tract into the cerebellar peduncle.

Added in the revised version: E, anti-peripherin-antibody (bright color) stained section of the elephant cerebellum. The cerebellar white matter is bright, the cerebellar granule cell layer has a light grey appearance, and the cerebellar molecular layer has a dark grey appearance. Accordingly, peripherin-reactivity mirrors the distribution of cerebellar climbing fibers (white matter > granule cell layers > molecular layer).

Added in the revised version: F, high magnification view of an anti-peripherin-antibody stained section of the elephant cerebellum. Climbing fibers are apparent by their elongated axonal pattern (arrow). Purkinje cells appear as bright dots. Even higher magnification revealed that the inside of Purkinje cells is dark (anti-peripherin-negative), i.e. the bright dot appearance of Purkinje cells reflects the ensheating of Purkinje cells by climbing fibers. wm = white matter; gcl = granule cell layer; pcl = Purkinje cell layer; ml = molecular layer.

Added in the revised version: A delineation of the olivo-cerebellar tract supports our partitioning scheme

Removed in the revised version: B, High magnification view of the dorsolateral nucleus (corresponding to the upper red rectangle in A). Anti-peripherin-positive axon bundles (putative climbing fibers) are seen in support of the inferior olive hypothesis of Reveyaz et al.

Removed in the revised version: C, High magnification view of the ventromedial nucleus (corresponding to the lower red rectangle in A). The ventromedial nucleus is weakly positive for peripherin but contains no anti-peripherin-positive axon bundles (i.e. no putative climbing fibers) in support of the trigeminal nucleus hypothesis of Reveyaz et al. Note that myelin stripes – weakly visible as dark omissions – are clearly anti-peripherin-negative.

Earlier version: Previous work on the elephant brainstem (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962) suggested that the structure we assigned as trigeminal nucleus is in fact the inferior olive; this matter is also discussed in depth in our earlier discussions with the referees, which are published along with our article. We performed additional peripherin-antibody staining to assess the possibility that the structure we assigned as trigeminal nucleus corresponds to the inferior olive. To differentiate between the trigeminal nuclei and the inferior olive, we used a climbing fiber antibody staining. Peripherin is a cytoskeletal protein that is found in peripheral nerves and climbing fibers. Specifically, climbing fibers of various species (mouse, rabbit, pig, cow, and human; Errante et al 1998) stain intensely with peripherin- antibodies. As can be seen in Figure 2, we observe peripherin-reactivity in axonal bundles (i.e. in putative climbing fibers), in what we think is the inferior olive (Figure 2B). We also observe some peripherin-reactivity in what we think is the trigeminal nucleus, but not the distinct and strong labeling of axonal bundles (Figure 2C). This lack of peripherin-reactive axon bundles suggests, that there are no climbing fibers, in what was previously thought as the inferior olive of the elephants. Peripherin-staining is consistent with the idea that the ventral nuclei described by us correspond to the trigeminal nuclei of the elephant and the more lateral nuclei, previously described as trigeminal nuclei by Maseko et al., 2013 correspond to the inferior olive.

Revised version: Previous work on the elephant brainstem (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962) suggested that the structure we assigned as trigeminal nucleus is in fact the inferior olive; this matter is also discussed in depth in our earlier discussions with the referees, which are published along with our article. We performed additional antibody staining to assess the possibility that the structure we assigned as trigeminal nucleus corresponds to the inferior olive. To differentiate between the trigeminal nuclei and the inferior olive, we used a climbing fiber antibody staining. Peripherin is a cytoskeletal protein that is found in peripheral nerves and climbing fibers. Specifically, climbing fibers of various species (mouse, rabbit, pig, cow, and human; Errante et al 1998) stain intensely with peripherin-antibodies. In Figure 2A we provide an overview of the elephant brainstem and key structures therein. Figure 2B shows schematically the key conclusions of our peripherin-antibody-staining, our delineation of the olivo-cerebellar tract, and indicates the position of peripherin-antibody-stained sections. We observed peripherin-reactivity in axonal bundles (i.e. in putative climbing fibers), in what we think is the inferior olive (Figure 2C, left and upper right, red stars). We also observe some peripherin-reactivity in what we think is the trigeminal nucleus, but not the distinct and strong labeling of axonal bundles (Figure 2C lower right, green stars). This lack of peripherin-reactive axon bundles suggests, that there are no climbing fibers, in what was previously thought of as the inferior olive of the elephants. We followed peripherin-reactive fibers through the brainstem and found they discharge into the cerebellar peduncle as expected for the olivo-cerebellar tract (Figure 2D, white arrow). Peripherin-reactivity was also observed in the cerebellum (Figure 2E), where putative climbing fibers ascend through the white matter and the granular cell layer and ensheath Purkinje cell somata (Figure 2F). These observations show that the strongly serrated dorsolateral nucleus connects to the cerebellum via climbing fibers. In contrast, the ventromedial nucleus, the putative trigeminal nucleus, does not receive climbing fibers. These data support our partitioning scheme over the assignments of Maseko et al 2013.

Trigeminal nuclei in coronal and horizontal sections of African elephant brainstem

We found that putatively trunk-related trigeminal modules repeat at least two and probably four times in the anterior-posterior direction in the elephant brainstem. All these repeating modules had a higher cytochrome oxidase reactivity and a higher cell density than surrounding brainstem structures. Such repeats of trigeminal representations in the anterior-posterior direction are also seen in other mammals (Belford & Killackey, 1979; Ma, 1991). We refer to these modules with the same terminology as established in rodents. As in other mammals, we found the most anterior representation to be larger than the others, and we refer to this representation as nucleus principalis (Pr5, which stands for principal trigeminal nucleus). Figure 3A shows a Nissl-stained coronal section through the principalis trigeminal modules. In Figure 3B, C we provide a color-coded putative assignment of principalis modules. We assigned the large (grey) module to the trunk, because of its cytochrome oxidase reactivity, its elongation, and its extraordinary size. We assigned the elongated (red/pink) module to the nostril for the following reasons: 1. its unusually (among brainstem modules) thin tube-like appearance. 2. The widening towards the tip of the putative trunk module. 3. The cellular continuity with the mouth opening of the putative trunk module. 4. The topographic relationship with the lower jaw module, which matched the topography of the elephant head (Figure 3C). 5. The fact that this module had the same length as the trunk module. 6. The fact that there were no indications of a nostril module inside the putative trunk module, (where we initially expected a nostril representation). Our reasons for assigning the compact (blue) module to the lower jaw were its shape (Figure 3A-C) and topographic position. Next, we provide an overview of the arrangement of trigeminal modules in horizontal sections (Figure 3D-F, proceeding from dorsal to ventral). At the dorsal level (Figure 3D) only two trunk modules (TM), can be recognized. These are Pr5 and Sp5o, which stands for spinal trigeminal nucleus pars oralis TM, directly posterior to the Pr5. The cell density is low, the modules barely stand out from the surroundings and we think that at the dorsal level proximal trunk parts are represented. At the midlevel (Figure 3E) four repeating trunk modules (Pr5 TM, Sp5o TM, and Sp5i, which stands for spinal trigeminal nucleus pars interpolaris, and Sp5c, which stands for spinal trigeminal nucleus pars caudalis TM) can be recognized. We did not investigate facial representations other than the trunk module. The identification of the Sp5i and Sp5c trunk modules is only tentative at this point. The analysis of horizontal and parasagittal sections pointed to a mirror image-like arrangement of these modules. The cell density is higher at midlevel (Figure 3E). At the ventral level (Figure 3F) only two trunk modules (Pr5 and Sp5o TM, directly posterior to the principalis) can be recognized. In this section, the mirror image-like arrangement of the Pr5 TM and the Sp5o TM is evident. The cell density is very high and we think that the trunk tip is represented here. We conclude that repeating trigeminal trunk modules can be recognized in the elephant brainstem.

Overview of trigeminal nuclei in coronal and horizontal sections of African elephant cow Indra

A, micrograph of Nissl-stained coronal 60-µm-section through the right-hemispheric principalis trunk module of elephant cow Indra. The principalis is the most anterior and by far the largest trigeminal representation in elephants.

B, drawing of color-coded trigeminal modules belonging to the principalis nucleus.

C, drawing of an African elephant head with different facial structures color-coded according to the trigeminal modules they putatively correspond to in A, B.

D, left micrograph of Nissl-stained horizontal section through the left-hemispheric trigeminal nuclei of elephant cow Indra. The section is positioned at a dorsal level of the trigeminal nuclei. Right, drawing of trigeminal modules shown in the left micrograph; for the principalis (Pr5) module the same color code as in B, and C has been used. Sp5o refers to the trunk module directly posterior to the principalis nucleus. The facial nucleus and the putative inferior olive were also identified.

E, conventions as in D. Horizontal section through the mid-level of the trigeminal nuclei. Sp5i TM and Sp5c TM, refer to putative trunk modules posterior to the Sp5o TM.

F, conventions as in D. Horizontal section through the ventral level of the trigeminal nuclei.

Pr5, principal trigeminal nucleus; Sp5o, spinal trigeminal nucleus pars oralis; Sp5i, spinal trigeminal nucleus pars interpolaris; Sp5c, spinal trigeminal nucleus pars caudalis; TM = trunk module; A = anterior; L = lateral; V = ventral.

The cellular architecture of the putative principalis trunk module

We studied the module’s cellular organization by Golgi stains of the putative principalis trunk module of the Asian elephant Raj (Figure 4A). Golgi stains identified two prominent neuronal elements. First, we observed bundles of large diameter (3-15 µm) axons, which run orthogonal to the module’s main axis in coronal sections (Figure 4B arrows). These axon bundles correspond to the myelin stripes described in Figure 5. Somato-dendritically stained neurons were the second neuronal element identified by Golgi stains (Figure 4C). The density of Golgi-stained neurons was very low. We reconstructed neurons using a Neurolucida system and superimposed 47 neurons (from three adjacent sections) in Figure 4D. The most abundant cells in the trigeminal nucleus were putative astrocytes (Figure 4E, green). We distinguished two types of neurons, putative principal neurons with large somata and branched dendrites (n = 41; black in Figure 4E) and other neurons, putative interneurons, with small somata and unbranched dendrites (n = 6; red in Figure 4E). Putative astrocytes, principal neurons, and interneurons differed markedly in their morphologies (Figure 4E, Table 1). The dendritic trees of neurons were elongated (Figure 4E), an observation confirmed when we prepared raw polar plots of dendritic orientation (Figure 4F, upper). We had the impression that dendritic elongation and axon bundles followed the same axis. We tested the idea that dendritic trees were aligned to myelin stripes by rotating dendritic trees and aligning all trees according to the local axon bundle orientations. When we aligned dendritic trees this way, we observed an even stronger population polarization of dendritic trees, i.e., dendritic trees were average twofold longer along the axon bundle axis (Figure 4F, lower).

Cellular organization of the putative trunk module in the elephant trigeminal nuclei.

A, Golgi stained 200 µm coronal section through the trunk module of Raj, a four-year-old Asian elephant bull.

B, bundles of very thick axons are revealed by the Golgi staining (red arrows). Bundles were oriented orthogonal to the main axis of the module except for the putative finger representation, where they curved around.

C, Golgi-stained neurons are also observed albeit at a low frequency.

D, well-stained and well-preserved neurons reconstructed with a Neurolucida system. We show 47 neurons reconstructed from three adjacent coronal Golgi sections superimposed. Note the small size of the neurons relative to the module.

E, left (green), ten putative astrocytes. These small cells are the most abundant cellular element in the trigeminal nucleus. Right, four neuronal reconstructions are shown at higher magnification. Putative principal cells (three shown in black, total n = 41) had large somata and branched dendrites. A few cells (one cell shown in red, total n = 6) had small somata and unbranched dendrites. Dendritic trees were weakly polarized.

F, upper, raw polar plot of the orientation of neuronal dendritic segments (from all putative principal cells and putative interneurons, n = 47) relative to the soma confirms the common elongation of dendrites. Lower, when cells were aligned to the local axon bundle orientation an even stronger polarization of dendrites is evident.

G, antibody staining of neurons (green fluorescence, NeuN-antibody) and nuclei of all cells (blue fluorescence, DAPI) of a coronal section through the putative trunk module of Indra, a 34-year-old female African elephant. Neuron density is low.

H, high magnification view of the section shown in G, non-neural cells outnumber neurons by about a hundredfold (data refer to neuron and non-neural cell counts from three elephant trigeminal nuclei).

I, upper, somata drawing from a Nissl stained 60 µm coronal section through the putative trunk module of Indra. Lower, cells from the medial, the putatively proximal trunk representation of the module, and the lateral, the putatively distal trunk representation. Note the soma size difference.

J, plot of soma area along the length of the module. Neurons were sequentially measured along the axis of the module. Each dot refers to one of 1159 neurons in the section; red running average (across 40 neurons) of soma area. Cells are significantly larger in lateral (putatively distal trunk representation), unpaired T-test.

FT, putative dorsal Finger Tip representation; V = ventral; L = lateral.

Trunk module myelin stripes form a precise map of trunk folds

A, a brightfield image of a freshly cut 60 µm coronal section through the center of the putative trunk module of adult female elephant Indra. Neurons are evident as small white dots and whitish myelin stripes are readily apparent even in this unstained tissue. B, a fluomyelin stain (green fluorescence) confirms that stripes contain myelin. High-resolution brightfield microscopy (not shown) and Golgi stains (Figure 2B) confirm that the stripes consist of axon bundles.

C, a line drawing (red) of myelin stripes superimposed to the micrograph shown in A.

D, upper, enlarged view of line drawing (red) of myelin stripes shown in C; we quantified dorsally ending, ventrally ending, and full transversal stripes. Such numbers match the number of trunk folds quantified in E.

Lower, based on the idea of a match of trunk folds with myelin stripes one can compute magnification factors across the trunk module. Neural data (green) refer to distances between dorsally ending myelin stripes (i.e., neural distances were measured along the dorsal border (green line) of the module). Trunk distances were measured between dorsal trunk folds.

E, drawing of the dorsal folds of the trunk of Indra.

F, the composite photograph of the dorsal trunk of Indra. We counted dorsal folds, ventral folds (not visible here), and folds that fully transversed the right trunk side (not visible here). Trunk folds match in number, orientation (typically transversal), and patterning with myelin stripes seen on the trunk module.

V = ventral; L = lateral.

We stained coronal sections with NeuN-antibody to identify neurons and used the DNA-stain DAPI to identify all cell nuclei (Figure 4G). Neuronal density was fairly low, but the density of non-neuronal cells was substantial (Figure 4H). In counts of individual fluorescence sections, we observed a ratio of neurons to non-neural cells of 1 to 80 (in one elephant). We then made more systematic counts of neurons vs. visually identified non-neural cells across three entire trunk modules from two elephants. We then observed a ratio of 1 neuron to 108 ± 24 non-neural cells. We also observed that neuron size was not homogeneous across the putative trunk module. We observed a cell size increase from proximal to distal in all putative modules (n = 12), for which we had cellular stains. To quantify cell size differences, we drew somata from a Nissl-stained section through the center of the trunk module of the African elephant cow Indra (Figure 4I). We found cell size to increase significantly from the putative proximal to the putative distal finger representation (Figure 4J). We conclude that the putative trunk module contains transversally running axon bundles and neurons, which are vastly outnumbered by non-neural cells.

Module myelin stripes match with number, orientation, and patterning of trunk folds

Earlier version: In Nissl or cytochrome-oxidase stains, we observed prominent myelin stripes apparent as white omissions. Remarkably, entirely unstained freshly-cut coronal brainstem sections showed the clearest stripe pattern in brightfield microscopy (Figure 5A). Fluorescent stains for myelin (fluomyelin) confirmed the presence of myelin (Figure 5B). As already suggested by Golgi staining, the myelin stripes appeared to consist of large-diameter axons. The visibility of myelin stripes varied with the sectioning plane and anterior-posterior position. Myelin stripes were most obvious at the anterior- posterior center of the putative trunk module, as seen for the coronal section in Figure 5A. We investigated the myelin-stripe trunk correspondence. To this end, we made drawings of myelin stripes (Figure 5C) and compared the pattern of myelin stripes (Figure 5D upper) to the pattern of trunk folds of Indra’s trunk (Figure 5E, F). Myelin stripe and trunk fold patterns were very similar. In all trunk modules sectioned, we observed an overall match of stripe and trunk fold orientation (to the module and trunk main axis, respectively). We also observed in all modules a lack of fully transversing stripes in the putative finger region of the module, which is consistent with the lack of folds across the trunk ‘mouth’. In favorable cases, where we had brightfield images of the trunk module and had access to the elephant’s trunk, the data hinted at a 1 to 1 matching of stripes and folds. Specifically, we observed 65 myelin stripes ending on the dorsal side of the module, 46 ventrally ending myelin stripes, and 31 full transversal myelin stripes. In terms of folds, we observed 64 dorsal trunk folds (Figure 5E), 49 folds on the ventral side of the trunk, and 32 folds that fully transversed the right side of Indra’s trunk. This numeric correspondence is very suggestive and inspired a detailed mapping of trunk sensory topography (Figure 5D lower). Based on stripe- wrinkle matching, we suggest that sensory magnification increases from 1000:1 (trunk: trigeminal nucleus) in the proximal representation of the trunk module to 5:1 in the trunk-finger representation. Sectioning angle was a major factor in determining the match between myelin stripes and folds, i.e. we observed myelin stripes in trunk fold- like patterns in all coronally sectioned specimens. In horizontally sectioned elephant brainstem, myelin stripes were seen but could not be related to folds. In parasagitally sectioned brainstems few myelin stripes were obvious. In coronal sections, myelin stripes were most obvious in the center of the module and matched best to trunk folds. The staining method was another determinant of the match. Myelin stripes were best visible in unstained freshly cut sections with brightfield microscopy. As expected, myelin stripes were also stained positively for fluorescent myelin dyes, such as fluomyelin-green (Figure 5B) or fluomyelin-red (data not shown). Nissl or cytochrome oxidase stains were less sensitive than visualizing myelin stripes in brightfield images, i.e. not all stripes are visible in each section. While in some sections like the one shown in Figure 5D, pretty much all stripes could be mapped to trunk folds, most sections contained a few stripes that had deviating trajectories from the other stripes and these stripes could not be mapped to trunk folds. A good match of myelin stripes to folds depended also on the assessment of trunk folds. Specifically, a good match was only obtained, if we restricted fold counts to major trunk wrinkles/folds, minor trunk wrinkles appear not to be robustly represented by myelin stripes. We conclude myelin stripe patterns behave not unlike rodent cortical barrel patterns, the visibility of which also greatly depends on the staining method and sectioning angle.

Revised version: In Nissl or cytochrome-oxidase stains, we observed prominent myelin stripes apparent as white omissions. Remarkably, entirely unstained freshly-cut coronal brainstem sections showed the clearest stripe pattern in brightfield microscopy (Figure 5A). Fluorescent stains for myelin (fluomyelin) confirmed the presence of myelin (Figure 5B). As already suggested by Golgi staining, the myelin stripes appeared to consist of large-diameter axons. The visibility of myelin stripes varied with the sectioning plane and anterior-posterior position. Myelin stripes were most obvious at the anterior-posterior center of the putative trunk module, as seen for the coronal section in Figure 5A. We investigated the myelin-stripe trunk correspondence. To this end, we made drawings of myelin stripes (Figure 5C) and compared the pattern of myelin stripes (Figure 5D upper) to the pattern of trunk folds of Indra’s trunk (Figure 5E, F). Myelin stripe and trunk fold patterns were very similar. In all trunk modules sectioned, we observed an overall match of stripe and trunk fold orientation (to the module and trunk main axis, respectively). We also observed in all modules a lack of fully transversing stripes in the putative finger region of the module, which is consistent with the lack of folds across the trunk ‘mouth’. In favorable cases, where we had brightfield images of the trunk module and had access to the elephant’s trunk, the data hinted at a 1 to 1 matching of stripes and folds. Specifically, we observed 65 myelin stripes ending on the dorsal side of the module, 46 ventrally ending myelin stripes, and 31 full transversal myelin stripes. In terms of folds, we observed 64 dorsal trunk folds (Figure 5E), 49 folds on the ventral side of the trunk, and 32 folds that fully transversed the right side of Indra’s trunk. This numeric correspondence is very suggestive and inspired a detailed mapping of trunk sensory topography (Figure 5D lower). Based on stripe-wrinkle matching, we suggest that sensory magnification increases from 1000:1 (trunk: trigeminal nucleus) in the proximal representation of the trunk module to 5:1 in the trunk-finger representation. Sectioning angle was a major factor in determining the match between myelin stripes and folds, i.e. we observed myelin stripes in trunk fold-like patterns in all coronally sectioned specimens. In horizontally sectioned elephant brainstem, myelin stripes were seen but could not be related to folds. In parasagitally sectioned brainstems few myelin stripes were obvious. In coronal sections, myelin stripes were most obvious in the center of the module and matched best to trunk folds.

Added in the revised version: The staining method was another determinant of the match. Myelin stripes were best visible in unstained freshly cut sections with brightfield microscopy. As expected, myelin stripes were also stained positively for fluorescent myelin dyes, such as fluomyelin-green (Figure 5B) or fluomyelin-red (data not shown). Nissl or cytochrome oxidase stains were less sensitive than visualizing myelin stripes in brightfield images, i.e. not all stripes are visible in each section. While in some sections like the one shown in Figure 5D, pretty much all stripes could be mapped to trunk folds, most sections contained a few stripes that had deviating trajectories from the other stripes and these stripes could not be mapped to trunk folds. A good match of myelin stripes to folds depended also on the assessment of trunk folds. Specifically, a good match was only obtained, if we restricted fold counts to major trunk wrinkles/folds, minor trunk wrinkles appear not to be robustly represented by myelin stripes. We conclude myelin stripe patterns behave not unlike rodent cortical barrel patterns, the visibility of which also greatly depends on the staining method and sectioning angle.

Myelin stripe architecture and the lack of a relation of stripes to trigeminal neurons

Their large size makes determining the architecture of myelin stripes challenging. To confront this challenge, we applied synchrotron-powered X-ray phase-contrast tomography of an 8 mm unstained and paraffin-embedded trigeminal nucleus tissue punch (Figure 6A, B); based on phase-contrast this methodology allows us to sample large image volumes (Figure 6C) with submicrometer (0.65 µm isotropic voxel size) resolution. Such imaging allowed us to identify myelin stripes in unstained trigeminal tissue (Figure 6D) and even enabled the reconstruction of individual large-diameter axons for several millimeters through the entire volume image (Figure 6E). As observed before with light microscopy, myelin stripes ran in the coronal plane and were about seven myelinated axons wide (maximum extent in the coronal plane; Figure 6F). Myelin stripes were circular axon bundles (Figure 6G) and were also about seven myelinated axons high (maximum extent in the anterior-posterior plane; Figure 6H). With that, we estimated that stripes are made up of 20-50 myelinated axons, unmyelinated axons could not be resolved in our analysis. The large-diameter axons, which could be followed through the X-ray tomography volume image followed an ‘all the way’ pattern (i.e. fully transversing the module). We also found that myelin stripes have a fairly consistent thickness from their dorsal to their ventral end (Figure 6I). This observation argues against the idea that myelin stripes are conventional axonal supply structures, from which axons divert off into the tissue. We also analyzed myelin stripes throughout the trunk module (Figure 6J) to understand how their thickness relates to trigeminal neuron numbers (i.e. the number of neurons between myelin stripes; Figure 6J). We observed little obvious relation between myelin stripe thickness and trigeminal neuron number. We conclude that myelin stripes have a stereotyped architecture, but show little relation to trigeminal neurons.

Microscopic organization of myelin stripes and absence of a strong relation of stripes to trigeminal neurons

A, upper, schematic of the trunk module with myelin stripes (grey) of African elephant Bambi and targeting of the 8 mm tissue punch (dashed red circle). Lower, synchrotron radiation (red flash, DESY, Hamburg) was directed to the area of the punch (black box) and imaged. V = ventral; L = lateral.

B, sketch of the parallel beam setup of the GINIX endstation (P10 beamline, DESY, Hamburg). In this geometry, a dataset of the trunk module was acquired at an effective voxel size of 650nm3

C, dimensions of the imaged volume shown as a volume rendering (0.65 µm isotropic voxel size)

D, transparent image volume. Two myelin stripes were followed through the volume image (highlighted red axon bundle). Four large diameter (∼ 15 µm) axons were also reconstructed and could be followed through the entire volume image (blue).

E, axon bundles (red) and reconstructed axons (blue) in isolation.

F, image section in the coronal plane at the center of the myelin stripe. The myelin stripe (pink overlay) is readily visible, the reconstructed axon is highlighted in blue, and the bundle has a width of about 7 myelinated axons.

G, cross-section through the axon bundle (pink overlay). The virtual section is cut orthogonal to the coronal plane. The reconstructed axon is highlighted in blue.

H, image section orthogonal to the coronal plane, the virtual section is cut in an anterior-posterior direction parallel to the axon bundle (pink overlay). The myelin stripe is readily visible, the reconstructed axon is highlighted in blue, and the bundle has a height of about 7 myelinated axons.

I, left, myelin stripes. Right, measurement of the thickness (width orthogonal to the main stripes axis) along the dorsoventral axis of stripes; data are the average of ten measurements along ten myelin stripes that fully transversed the module. Stripe thickness is fairly constant, there is no evidence of stripe tapering as would be expected if axons bud off into the tissue.

J, upper, measurement of the thickness (width orthogonal to the main stripes axis) of myelin stripes across the putative trunk module; only full transversal stripes were measured. Stripe thickness is fairly constant. Lower, neuron number between full transversal myelin stripes. Neuron number varies more than 100-fold and is very low (zero) between medial stripes (in the putatively proximal trunk representation). The fact that stripe thickness changes little across the module, while neuron number between stripes changes massively argues against a relationship between stripe thickness and trigeminal neuron number.

The putative trunk module mirrors species differences in trunk folds and trunk use

Earlier version: We found that the trigeminal bumps on the ventral brainstem differ significantly between African and Asian elephants (Figure 7A). We counted neurons in the principalis trunk module and found that African elephants (740210 ± 51902, mean ± SD) had more neurons than Asian elephants (636447 ± 69729, mean ± SD) and also had a larger volume principalis trunk module (Figure 7B). Supplementary Tables 2 and 3 provide further information on our counts of the trigeminal nuclei. As noted there, the dorsal finger accounted for a large fraction (∼20%) of the trunk modules. We wondered, why brainstem bumps differed between African and Asian elephants, and therefore closely investigated the shape of trunk modules in these species. A cytochrome-oxidase stained coronal section through the trunk module of the African elephant Indra is shown in Figure 7C. Drawings of coronal sections from this trunk module and the trunk module of other African elephants are shown in Figure 7D; myelin stripes (violet) were visible as whitish omissions of the cytochrome oxidase or of Nissl stains. We also determined the length, the width, and the longitudinal position of the greatest width of the module (black line; Figure 7D). A cytochrome-oxidase- stained coronal section through the trunk module of Asian elephant Raj is shown in Figure 7E and drawings from this and other trunk modules of Asian elephants are shown in Figure 7F. Basic aspects of the module trunk were similar in African and Asian elephants (Figure 7C-F), but the details differed. First, the African elephant trunk modules had fewer and thicker myelin stripes (Figure 7D, F); this is a most interesting observation since African elephants have fewer trunk folds than Asian elephants.

Revised version: We found that the trigeminal bumps on the ventral brainstem differ significantly between African and Asian elephants (Figure 7A). We counted neurons in the principalis trunk module and found that African elephants (740210 ± 51902, mean ± SD) had more neurons than Asian elephants (636447 ± 69729, mean ± SD) and also had a larger volume principalis trunk module (Figure 7B). Supplementary Tables 2 and 3 provide further information on our counts of the trigeminal nuclei. As noted there, the dorsal finger accounted for a large fraction (∼20%) of the trunk modules. We wondered, why brainstem bumps differed between African and Asian elephants, and therefore closely investigated the shape of trunk modules in these species. A cytochrome-oxidase stained coronal section through the trunk module of the African elephant Indra is shown in Figure 7C. Drawings of coronal sections from this trunk module and the trunk module of other African elephants are shown in Figure 7D; myelin stripes (violet) were visible as whitish omissions of the cytochrome oxidase or of Nissl stains. We also determined the length, the width, and the longitudinal position of the greatest width of the module (black line; Figure 7D). A cytochrome-oxidase-stained coronal section through the trunk module of Asian elephant Raj is shown in Figure 7E and drawings from this and other trunk modules of Asian elephants are shown in Figure 7F. Basic aspects of the module trunk were similar in African and Asian elephants (Figure 7C-F), but the details differed. First, the African elephant trunk modules had fewer and thicker myelin stripes (Figure 7D, F); this is a most interesting observation since African elephants have fewer trunk folds than Asian elephants. Asian elephants have more of these folds, but they are also more shallow and less pronounced (Schultz et al., 2023). Given our limited material and that myelin stripes are less conspicuous in Asian elephants than in African elephants, we were not able to ascertain the one-to-one match of stripes and folds that we could show for the African elephant Indra in Figure 5. Second, African elephant trunk modules were significantly longer but not wider (Figure 7G). Asian elephant trunk modules had a much more round-shaped appearance than African elephant trunk modules. The greatest width of the Asian elephant trunk module is at positions representing the trunk wrapping zone, which we determined from photographs. African elephant trunk modules were also significantly more ‘top-heavy’, i.e. they had their widest point much closer to the putative trunk tip (Figure 7H). We suggest that the shape differences between African and Asian elephant trunk modules might be related to the different grasping strategies of these two elephant species (Racine, 1980). African elephants have two fingers and tend to pinch objects (Figure 7I; upper), a grasping strategy that emphasizes the trunk tip in line with their ‘top-heavy’ trunk module. Asian elephants in contrast have only one finger and tend to wrap objects with their trunk (Racine, 1980; Figure 5I; lower). This grasping strategy engages more of the trunk and, in line with this behavior, the width of the Asian elephant trigeminal nucleus is maximal in the trunk wrapping area.

Differences between the putative trigeminal trunk modules of Asian (Elephas maximus) and African (Loxodonta africana) elephants

A, left, ventral view of the brainstem of African elephant Indra. Right, ventral view of the brainstem of Asian elephant Dumba. The anterior-posterior length of the trigeminal nuclei from the pons is indicated as a black line. The average trigeminal nuclei length refers to 10 African and 10 Asian trigeminal nuclei, the p-value refers to a Mann-Whitney test. The trigeminal nuclei bump is more elongated in African than in Asian elephants.

B, neuron number (left) and volume (right) of the putative principalis trunk module in Asian and African elephants. Trigeminal nuclei come from three African and three Asian elephants; p-values refer to unpaired t-tests.

C, micrograph of a cytochrome-oxidase-stained section through the putative trunk module of African elephant Indra.

D, drawings of the outlines and myelin stripes from cytochrome oxidase or Nissl stained sections through the putative trunk module of African elephants; the top drawing was made from the micrograph shown in C. The black line refers to the point of greatest width along the direction of myelin stripes on the putative trunk shaft (the putative finger was not considered in the width analysis).

E, micrograph of a cytochrome-oxidase-stained section through the putative trunk module of Asian elephant Raj.

F, drawings, of putative trunk modules from Asian elephants; conventions as in D.

G, length, and width of putative trunk module in African and Asian elephants. p-values refer to t-tests.

H, upper, drawing of the trunk of an African elephant. Lower, drawing of the trunk of an Asian elephant; note that Asian elephants have more folds. Arrows mark the point of greatest width of the putative trunk module (black lines in D, F) projected back on trunk positions in African (upper) and Asian (lower) elephant trunks. We also highlighted in color the dorsal (red) and ventral (pink) trunk tip and wrapping zone of Asian elephants in green and the analogous trunk part of African elephants in light green. The extent of the trunk wrapping zone was determined from photographs of Asian elephants wrapping objects. Specifically, we defined the wrapping zone as the trunk parts in contact with large objects (mangos, melons, fodder beets) being wrapped.

I, object grasping/pinching behavior in African (upper) and object wrapping strategy in Asian (lower) elephants (adapted from Kaufmann et al., 2022).

Removed in the revised version: Asian elephants have more of these folds, but they are also more shallow and less pronounced (Schultz et al., 2023). Given our limited material and that myelin stripes are less conspicuous in Asian elephants than in African elephants, we were not able to ascertain the one-to-one match of stripes and folds that we could show for the African elephant Indra in Figure 5. Second, African elephant trunk modules were significantly longer but not wider (Figure 7G). Asian elephant trunk modules had a much more round-shaped appearance than African elephant trunk modules. The greatest width of the Asian elephant trunk module is at positions representing the trunk wrapping zone, which we determined from photographs. African elephant trunk modules were also significantly more ‘top-heavy’, i.e. they had their widest point much closer to the putative trunk tip (Figure 7H). We suggest that the shape differences between African and Asian elephant trunk modules might be related to the different grasping strategies of these two elephant species (Racine, 1980). African elephants have two fingers and tend to pinch objects (Figure 7I; upper), a grasping strategy that emphasizes the trunk tip in line with their ‘top-heavy’ trunk module. Asian elephants in contrast have only one finger and tend to wrap objects with their trunk (Racine, 1980; Figure 5I; lower). This grasping strategy engages more of the trunk and, in line with this behavior, the width of the Asian elephant trigeminal nucleus is maximal in the trunk wrapping area.

Discussion

Summary

We describe a pair of large bumps on the ventral surface of the elephant medulla that contain metabolically highly active, densely vascularized repeating modules. The trunk module contains an accurate myelin map of trunk folds. Mapping myelin stripes to the trunk folds indicated an increase in sensory magnification from the proximal to the distal trunk. Magnification analysis also identified an enlarged trunk wrapping zone in Asian elephants, who wrap objects with their trunk.

The ventral brainstem bumps likely correspond to elephant trigeminal nuclei

Added in the revised version: Establishing elephant brainstem organization is challenging because both tracing methods and in vivo electrophysiology cannot be applied to elephants. Our assignments of trigeminal nuclei deviate from earlier suggestions (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962), which assigned the putative trigeminal nuclei as inferior olive, and the structure identified as inferior olive, as trigeminal nuclei. For several reasons, we think that our partitioning scheme with ventromedial trigeminal nuclei and a dorsolateral inferior olive is superior to the scheme of Maseko et al 2013 with a ventromedial inferior olive and dorsolateral trigeminal nuclei. Our synopsis of the evidence is the following. First of all, we agree with that concerning brainstem position our scheme of a ventromedial trigeminal nucleus and a dorsolateral inferior olive deviates from the usual mammalian position of these nuclei (i.e. a dorsolateral trigeminal nucleus and a ventromedial inferior olive). However, cytoarchitectonics support our partitioning scheme. The compact cellular appearance of our ventromedial trigeminal nucleus is characteristic of trigeminal nuclei. The serrated appearance of our dorsolateral inferior olive is characteristic of the mammalian inferior olive. To our knowledge, nobody has described a mammalian trigeminal nucleus with a serrated appearance, that would apply to the elephant according to Maseko et al. 2013. Furthermore, metabolic staining (cyto-chrome-oxidase reactivity) supports our partitioning scheme. Specifically, our ventromedial trigeminal nucleus shows intense cyto-chrome-oxidase reactivity as it is seen in the trigeminal nuclei of trigeminal tactile experts. Additionally, the myelin stripes on the ventromedial trigeminal nucleus are isomorphic to trunk wrinkles. Isomorphism is a characteristic of somatosensory brain structures (barrel, barrelettes, nose-stripes, etc) and we know of no case, where such isomorphism was misleading. The large-scale organization of the ventromedial trigeminal nuclei in anterior-posterior repeats is characteristic of the mammalian trigeminal nuclei. To our knowledge, no such organization has ever been reported for the inferior olive. Finally, the connectivity analysis supports our partitioning scheme. According to our delineation of the elephant olivo-cerebellar tract, our dorsolateral inferior olive is connected via peripherin-positive climbing fibers to the cerebellum. In contrast, our ventromedial trigeminal nucleus (the referee’s inferior olive) is not connected via climbing fibers to the cerebellum. The above arguments suggest that the elephant trigeminal nuclei very likely correspond to the ventral brainstem bumps (please also see our commentary to referees published along with the article). As clearly expressed by the comments of Referee 2 published along with our article, the strength of other partitioning schemes, which assign the ventral brainstem bump as inferior olive is the better positional match of the assignment with the inferior olive of other mammals. We think this is a valid criticism. Our model offers coherent anatomical entities with trigeminal nuclei that look like trigeminal nuclei of other mammals, a trunk module with a striking resemblance to the trunk, and an inferior olive that looks like the inferior olive of other mammals.

Removed in the revised version: Establishing elephant brainstem organization is challenging because both tracing methods and in vivo electrophysiology cannot be applied to elephants. Our assignments of trigeminal nuclei deviate from earlier suggestions (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962), which assigned the putative trigeminal nuclei as inferior olive, and the structure identified as inferior olive, as trigeminal nuclei. For several reasons we think the ventral brainstem bumps correspond to the elephant trigeminal nuclei: 1. Their intense cytochrome oxidase reactivity, which is characteristic of the trigeminal nuclei of tactile specialists (Belford & Killackey, 1979; Ma, 1991; Catania bet al. 2013; Sawyer et al. 2014; Sawyer et al. 2015). 2. The large size of the putative trunk module, matches with the large infraorbital nerve of elephants (Purkart et al., 2022). 3. The elongation of the putative trunk module, which aligns with the trunk shape. 4. The arrangement of putative trigeminal modules (trunk module, nostril module, lower jaw module) that matches with elephant head topography. 5. The matching of module myelin stripes to trunk folds. 6. The organization of the trigeminal nuclei in modules repeats in the anterior-posterior direction, similar to other mammals (Belford & Killackey, 1979; Ma, 1991).

Removed in the revised version: The idea that ventral brainstem bumps correspond to inferior olive and that the laterally placed nuclei (referred to as inferior olive by us) correspond to the trigeminal nuclei (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962) appears unlikely: 1. The lateral nuclei (referred to as inferior olive by us) have the characteristic laminar cellular arrangement seen in the inferior olive of all mammals. 2. There are no indications of trigeminal-like modules in the lateral nuclei. 3. The lateral nuclei are not organized in modules repeating in the anterior-posterior direction, an arrangement seen in the trigeminal nuclei of other mammals 4. The ventral brainstem bumps (referred to as trigeminal nuclei by us) do not have the characteristic laminar cellular arrangement seen in the inferior olive of all mammals. 5. The very intense staining for cytochrome oxidase reactivity, seen in the ventral brainstem bumps, is not characteristic of the mammalian inferior olive. 6. The absence of climbing fiber bundles in the ventral brainstem bumps (Figure 2), which lets it appear unlikely that these bumps correspond to the inferior olive.

Removed in the revised version: A synopsis of the above arguments suggests that the elephant trigeminal nuclei very likely correspond to the ventral brainstem bumps (please also see our Referee Figure 3 and Referee Table 1 in our response letter published along with the article). As clearly expressed by the comments of Referee 2 published along with our article, the strength other partitioning scheme, which assign the ventral brainstem bump as inferior olive is the better positional match of the assignment with the inferior olive of other mammals. We think this is a valid criticism. Our model offers single anatomical nuclei with trigeminal nuclei that look like trigeminal nuclei of other mammals, a trunk module with a striking resemblance to the trunk, and an inferior olive that looks like the inferior olive of other mammals.

Detailed neuroanatomic mapping of elephant trigeminal nuclei

Earlier version: Our work provided detailed mapping of the elephant trigeminal brainstem into four repeating nuclei, consisting of several facial modules (most prominently the trunk module, the nostril module, and the lower jaw module). Because of the level of detail of our topography suggestions, however, we think it will be relatively straightforward to test the validity of our topography suggestions. Specifically, we predict that the dorsal trunk finger representation of the principalis trunk module will be connected with the distal part of the dorsal subnucleus of the facial nucleus (which contains the putative motor representation of the dorsal trunk finger (Kaufmann et al., 2022)). We would also predict that the dorsal trunk finger representation of the principalis trunk module will be connected with the dorsal trunk finger representation of the oralis nucleus. These connectivity suggestions are in the 1-2 cm range and can be tested with postmortem tracers like DiI. We also predict fewer myelin stripes in trunk modules of elephants with particularly few trunk folds (newborn or fetuses of African elephants) compared to adult African elephants.

Revised version: The protruding of the ventral brainstem bumps in the elephant brain reminds us of the protruding of layer 2 cell clusters in the human entorhinal cortex (Solodkin and Van Hoesen 1996). More generally speaking myelin stripes of the putative elephant trigeminal nuclei, are another example of isomorphic representation in the somatosensory system. Other classic examples of such isomorphic representation are cortical barrels (Woolsey and Van der Loos), brainstem barrelettes (Belford & Killackey, 1979) in the whisker system. Also nose-related isomorphisms have been described before, i.e. the pig cortical rostrum gyrus (Ritter al. 2021), or the stripe-like representation of nose appendages in the brainstem of the star-nosed mole (Catania, Leitch and Gauthier 2011). Our work provided detailed mapping of the elephant trigeminal brainstem into four repeating nuclei, consisting of several facial modules (most prominently the trunk module, the nostril module, and the lower jaw module). Because of the level of detail of our topography suggestions, however, we think it will be relatively straightforward to test the validity of our topography suggestions. Specifically, we predict that the dorsal trunk finger representation of the principalis trunk module will be connected with the distal part of the dorsal subnucleus of the facial nucleus (which contains the putative motor representation of the dorsal trunk finger (Kaufmann et al., 2022)). We would also predict that the dorsal trunk finger representation of the principalis trunk module will be connected with the dorsal trunk finger representation of the oralis nucleus. These connectivity suggestions are in the 1-2 cm range and can be tested with postmortem tracers like DiI. We also predict fewer myelin stripes in trunk modules of elephants with particularly few trunk folds (newborn or fetuses of African elephants) compared to adult African elephants.

A myelin map of trunk folds, white matter function and myeloarchitecture

Earlier version: According to conventional wisdom, neurons (gray matter) are the site of processing, and myelinated axons (white matter) are a subjugated supply system, whose sole function is to bring the correct axonal input to neurons. Our observations on trunk module myelin stripes are at odds with this view of myelin. Specifically, myelin stripes show no tapering (which we would expect if axons divert off into the tissue). More than that, there is no correlation between myelin stripe thickness (which presumably correlates with axon numbers) and trigeminal module neuron numbers. Thus, there are numerous myelinated axons, where we observe few or no trigeminal neurons. These observations are incompatible with the idea that myelin stripes form an axonal ‘supply’ system or that their prime function is to connect neurons. What do myelin stripe axons do, if they do not connect neurons? We suggest that myelin stripes serve to separate rather than connect neurons. Specifically, trunk module myelin stripes look like a map of trunk folds. Myelin stripes match with the number, orientation, and species-specific patterning of trunk folds. We note that if myelin stripes would behave as an axonal ‘supply’ system, they would be very thin/invisible in the proximal trunk, proximal trunk folds would not be visible, and distal stripes should be very thick. If myelinated axons have a life of their own and do not simply go where they find target neurons, we need to analyze them in novel ways. In particular, it seems to be a good idea to ‘look’ at patterns of myelination, rather than to immediately assume that this is a supply/connectivity system. We note early neuroanatomists like Oskar Vogt (Vogt, 1911; Niewenhuys, Broere & Cerliani, 2015) described incredibly intricate patterns of intracortical myeloarchitecture, patterns that are not easily explained in terms of a connectivity system to this day. In conclusion, we propose a novel white-matter function, which is to separate and functionally demarcate neurons as opposed to the conventionally assumed white-matter function of connecting neurons.

Revised version: According to conventional wisdom, neurons (gray matter) are the site of processing, and myelinated axons (white matter) are a subjugated supply system, whose sole function is to bring the correct axonal input to neurons. Our observations on trunk module myelin stripes are at odds with this view of myelin. Specifically, myelin stripes show no tapering (which we would expect if axons divert off into the tissue). More than that, there is no correlation between myelin stripe thickness (which presumably correlates with axon numbers) and trigeminal module neuron numbers. Thus, there are numerous myelinated axons, where we observe few or no trigeminal neurons. These observations are incompatible with the idea that myelin stripes form an axonal ‘supply’ system or that their prime function is to connect neurons. What do myelin stripe axons do, if they do not connect neurons? We suggest that myelin stripes serve to separate rather than connect neurons. Specifically, trunk module myelin stripes look like a map of trunk folds. Myelin stripes match with the number, orientation, and species-specific patterning of trunk folds. We note that if myelin stripes would behave as an axonal ‘supply’ system, they would be very thin/invisible in the proximal trunk, proximal trunk folds would not be visible, and distal stripes should be very thick. If myelinated axons have a life of their own and do not simply go where they find target neurons, we need to analyze them in novel ways. In particular, it seems to be a good idea to ‘look’ at patterns of myelination, rather than to immediately assume that this is a supply/connectivity system. We note early neuroanatomists like Oskar Vogt (Vogt, 1911; Niewenhuys, Broere & Cerliani, 2015) described incredibly intricate patterns of intracortical myeloarchitecture, patterns that are not easily explained in terms of a connectivity system to this day. We reckon that the exciting novel methodologies for determining myelo-architecture (Haenelt et al. 2023) will bring the issues and unanswered questions raised here to the foreground of neuroscientific inquiry. In conclusion, we propose a novel white-matter function, which is to separate and functionally demarcate neurons as opposed to the conventionally assumed white-matter function of connecting neurons.

Trigeminal organization in Asian and African elephants

Earlier version: At first sight, Asian and African elephant trigeminal nuclei are very similar. Both elephants have big ventral brainstem bumps, which contain the same modules (a putative trunk, nostril, and lower jaw module) and the nuclei stain intensely for cytochrome oxidase reactivity. A closer look reveals species differences, however, which may relate to the different trunk grasping strategies of Asian and African elephants. The first difference refers to the shape of the ventral brainstem bump, which is more roundish and shorter in Asian elephants and more elongated in African elephants. To our knowledge, this ventral brainstem bump is the only hitherto described difference, which allows us to differentiate Asian and African elephant brains from the outside. The length difference between Asian and African elephant ventral bumps reflects the different shapes of Asian and African elephant trunk modules. The African elephant trunk module is notably longer, more slender, and more top-heavy than the swaged Asian elephant trunk module. As we pointed out in Figure 5, such differences imply an enlargement of the trunk wrapping zone in the Asian elephant trunk module, in line with the object-wrapping behavior of Asian elephants (Racine 1980). Similarly, the top-heavy shape of the African elephant trunk module could be instrumental in the object-pinching of African elephants (Racine, 1980). These trigeminal differences are reminiscent of similar Asian-African species differences in the elephant facial nucleus (Kaufmann et al., 2022). We conclude that grasping behavior shapes the species-specific architecture of the trigeminal nuclei.

Revised version: At first sight, Asian and African elephant trigeminal nuclei are very similar. Both elephants have big ventral brainstem bumps, which contain the same modules (a putative trunk, nostril, and lower jaw module) and the nuclei stain intensely for cytochrome oxidase reactivity. A closer look reveals species differences, however, which may relate to the different trunk grasping strategies of Asian and African elephants. The first difference refers to the shape of the ventral brainstem bump, which is more roundish and shorter in Asian elephants and more elongated in African elephants. To our knowledge, this ventral brainstem bump is the only hitherto described difference, which allows us to differentiate Asian and African elephant brains from the outside. The length difference between Asian and African elephant ventral bumps reflects the different shapes of Asian and African elephant trunk modules. The African elephant trunk module is notably longer, more slender, and more top-heavy than the swaged Asian elephant trunk module. As we pointed out in Figure 7, such differences imply an enlargement of the trunk wrapping zone in the Asian elephant trunk module, in line with the object-wrapping behavior of Asian elephants (Racine 1980). Similarly, the top-heavy shape of the African elephant trunk module could be instrumental in the object-pinching of African elephants (Racine, 1980). These trigeminal differences are reminiscent of similar Asian-African species differences in the elephant facial nucleus (Kaufmann et al., 2022). We conclude that grasping behavior shapes the species-specific architecture of the trigeminal nuclei.

Conclusion

The elephant brainstem is exquisitely well-ordered and contains very large and detailed trigeminal representations. Trunk module myelin stripes form a map of trunk folds and accordingly serve to functionally separate neurons rather than to connect them. Further work should test the predictions of sensory topographies outlined here and ask, what further insights the elephant brain provides about the organization of gray and white matter.